An arms race between 5’ppp-RNA virus and its alternative recognition receptor MDA5 in RIG-I-lost teleost fish

Author response:

The following is the authors’ response to the previous reviews.

eLife assessment

The authors present evidence suggesting that MDA5 can substitute as a sensor for triphosphate RNA in a species that naturally lacks RIG-I. The key findings are potentially important for our understanding of the evolution of innate immune responses. Compared to an earlier version of the paper, the strength of evidence has improved but it is still partially incomplete due to a few key missing experiments and controls.

We would like to thank the editorial team for their positive comments and constructive suggestions on improving our manuscript. We have made further improvements based on the valuable suggestions of the reviewers, and we are pleased to send you the revised manuscript now. After revising the manuscript and further supplementing with experiments, we think that our existing data can support our claims.

Public Reviews:

Reviewer #1 (Public Review):

This study offers valuable insights into host-virus interactions, emphasizing the adaptability of the immune system. Readers should recognize the significance of MDA5 in potentially replacing RIG-I and the adversarial strategy employed by 5'ppp-RNA SCRV in degrading MDA5 mediated by m6A modification in different species, further indicating that m6A is a conservational process in the antiviral immune response.

However, caution is warranted in extrapolating these findings universally, given the dynamic nature of host-virus dynamics. The study provides a snapshot into the complexity of these interactions, but further research is needed to validate and extend these insights, considering potential variations across viral species and environmental contexts. Additionally, it is noted that the main claims put forth in the manuscript are only partially supported by the data presented.

After meticulous revisions of the manuscript, including adjustments to the title, abstract, results, and discussion, the main claim of our study now is the arm race between the MDA5 receptor and SCRV virus in a lower vertebrate fish, M. miiuy. This mainly includes two parts: Firstly, the MDA5 of M. miiuy can recognize virus invasion and initiate host immune response by recognizing the triphosphate structure of SCRV. Secondly, as an adversarial strategy, 5’ppp-RNA SCRV virus can utilize the m6A mechanism to degrade MDA5 in M. miiuy. Based on the reviewer's suggestions, we have further supplemented the critical experiments (Figure 3F-3G, Figure 4D, Figure 5G) and provided a more detailed and accurate explanation of the experimental conclusions, we believe that our existing manuscript can support our main claims. In addition, because virus-host coevolution complicates the derivation of universal conclusions, we will further expand our insights in future research.

Reviewer #2 (Public Review):

This manuscript by Geng et al. aims to demonstrate that MDA5 compensates for the loss of RIG-I in certain species, such as teleost fish miiuy croaker. The authors use siniperca cheats rhabdovirus (SCRV) and poly(I:C) to demonstrate that these RNA ligands induce an IFN response in an MDA5-dependent manner in m.miiuy derived cells. Furthermore, they show that MDA5 requires its RD domain to directly bind to SCRV RNA and to induce an IFN response. They use in vitro synthesized RNA with a 5'triphosphate (or lacking a 5'triphosphate as a control) to demonstrate that MDA5 can directly bind to 5'-triphosphorylated RNA. The second part of the paper is devoted to m6A modification of MDA5 transcripts by SCRV as an immune evasion strategy. The authors demonstrate that the modification of MDA5 with m6A is increased upon infection and that this causes increased decay of MDA5 and consequently a decreased IFN response.

One critical caveat in this study is that it does not address whether ppp-SCRV RNA induces IRF3-dimerization and type I IFN induction in an MDA5 dependent manner. The data demonstrate that mmiMDA5 can bind to triphosphorylated RNA (Fig. 4D). In addition, triphosphorylated RNA can dimerize IRF3 (4C). However, a key experiment that ties these two observations together is missing.

Specifically, although Fig. 4C demonstrates that 5'ppp-SCRV RNA induces dimerization (unlike its dephosphorylated or capped derivatives), this does not proof that this happens in an MDA5-dependent manner. This experiment should have been done in WT and siMDA5 MKC cells side-by-side to demonstrate that the IRF3 dimerization that is observed here is mediated by MDA5 and not by another (unknown) protein. The same holds true for Fig. 4J.

Thank you for the referee's professional suggestions. In fact, we have transfected SCRV RNA into WT and si-MDA5 MKC cells, and subsequently assessed the dimerization of IRF3 and the IFN response (Figure 2P-2Q). The results indicated that knockdown of MDA5 prevents immune activation of SCRV RNA. However, considering the potential for SCRV RNA to activate immunity independent of the triphosphate structure, this experimental observation does not comprehensively establish the MDA5-dependent induction of IRF3 dimer by 5’ppp-RNA. Accordingly, in accordance with the referee's recommendation, we proceeded to investigate the inducible activity of 5'ppp-SCRV on IRF3 dimerization in WT and si-MDA5 MKC cells, revealing that 5'ppp-SCRV indeed elicits immunity in an MDA5-dependent manner (Figure 4D). Additionally, poly(I:C)-HMW, a known ligand for MDA5, demonstrated a residual, albeit attenuated, activation of IRF3 following MDA5 knockdown, potentially attributed to its capacity to stimulate immunity through alternative pathways such as TLR3.

- Fig 1C-D: these experiments are not sufficiently convincing, i.e. the difference in IRF3 dimerization between VSV-RNA and VSV-RNA+CIAP transfection is minimal.

We have reconstituted the necessary materials and repeated the pertinent experiments depicted in Fig 1C-1D. The results demonstrate that SCRV-RNA+CIAP and VSV-RNA+CIAP exhibit a mitigating effect on the induction activity of SCRV-RNA and VSV-RNA on IRF3 dimerization, albeit without complete elimination (Figure 1C and 1D). These findings suggest the presence of receptors within M. miiuy and G. gallus capable of recognizing the viral triphosphate structure; however, it is worth noting that RNA derived from SCRV and VSV viruses does not exclusively depend on the triphosphate structure to activate the host's antiviral response.

Fig. 2N and 2O: why did the authors decide to use overexpression of MDA5 to assess the impact of STING on MDA5-mediated IFN induction? This should have been done in cells transfected with SCRV or polyIC (as in 2D-G) or in infected cells (as in 2H-K). In addition, it is a pity that the authors did not include an siMAVS condition alongside siSTING, to investigate the relative contribution of MAVS versus STING to the MDA5-mediated IFN response. Panel O suggests that the IFN response is completely dependent on STING, which is hard to envision.

In our previous laboratory investigations, we have substantiated the induction effect of STING on IFN under SCRV infection or poly(I:C) stimulation, as documented in the relevant literature (10.1007/s11427-020-1789-5), which we have referenced in our manuscript (lines 177-178). While we did assess the impact of STING on MDA5-mediated IFN induction in SCRV-infected cells, as indicated in the figure legends, we have revised Figure 2N-2O for improved clarity, and similarly, Figure 1H-1I has also been updated. Furthermore, considering that RNA virus infection can activate the cGAS/STING axis (10.3389/fcimb.2023.1172739) and the significant role of MAVS in sensing RNA virus invasion in the NLR pathway (10.1038/ni.1782), it is challenging to ascertain the respective contributions of STING and MAVS to the immune signaling cascade mediated by MDA5 during RNA virus infection. We intend to explore this aspect further in future research endeavors.

Fig. 3F and 3G: where are the mock-transfected/infected conditions? Given that ectopic expression of hMDA5 is known to cause autoactivation of the IFN pathway, the baseline ISG levels should be shown (ie. In absence of a stimulus or infection). Normalization of the data does not reveal whether this is the case and is therefore misleading.

Based on the reviewer's suggestions, we have rerun the experiment. We examined the effects of MDA5 and MDA5-ΔRD on antiviral factors in both uninfected, SCRV-infected, and poly(I:C)-HMW-stimulated MKC cells. Results showed that overexpression of both MDA5 and MDA5-ΔRD stimulated the expression of antiviral genes. However, when cells were infected or stimulated with SCRV or poly(I:C)-HMW, only the overexpression of MDA5, not MDA5-ΔRD, significantly increased the expression of antiviral genes (Figure 3F-3I).

Fig. 4F and 4G: can the authors please indicate in the figure which area of the gel is relevant here? The band that runs halfway the gel? If so, the effects described in the text are not supported by the data (i.e. the 5'OH-SCRV and 5'pppGG-SCRV appear to compete with Bio-5'ppp-SCRV as well as 5'ppp-SCRV).

Apologies for any confusion. The relevant areas in the gel pertaining to the experimental findings were denoted with asterisks and elaborated upon in the figure legends (Figure 4G, 4H, and 4M). The findings indicated that 5'ppp-SCRV, in contrast to 5'OH-SCRV and 5'pppGG-SCRV, demonstrated the ability to compete with bio-5'ppp-SCRV.

My concerns about Fig. 5 remain unaltered. The fact that MDA5 is an ISG explains its increased expression and increased methylation pattern. The authors should at the very least mention in their text that MDA5 is an ISG and that their observations may be partially explained by this fact.

First, as our m6A change analysis pipeline controls for changes in gene expression, these data should represent true changes in m6A modification rather than changes in the expression of m6A-modified transcripts (10.1038/s41598-020-63355-3). Similar studies demonstrated that m6A modification in RIOK3 and CIRBP mRNAs are altered following Flaviviridae infection (10.1016/j.molcel.2019.11.007). The specific calculation method is as follows: relative m6A level for each transcript was calculated as the percent of input in each condition normalized to that of the respective positive control spike-in. Fold change of enrichment was calculated with mock samples normalized to 1. Therefore, changes in the expression level of MDA5 can partially explain the increase in m6A modification on all MDA5 mRNA in cells, but it cannot indicate changes in m6A modification on each mDA5 transcript. We have supplemented the calculation method process in the manuscript and cited relevant literature (Lines 606-608). In addition, we have elaborated on the fact that MDA5 is an ISG gene in the experimental results (lines 260-261), and emphasized its compatibility with enhanced m6A modification of MDA5 in the discussion section (lines 405-409).

Reviewer #3 (Public Review):

In this manuscript, the authors explored the interaction between the pattern recognition receptor MDA5 and 5'ppp-RNA in the Miiuy croaker. They found that MDA5 can serve as a substitute for RIG-I in detecting 5'ppp-RNA of Siniperca cheilinus rhabdovirus (SCRV) when RIG-I is absent in Miiuy croaker. Furthermore, they observed MDA5's recognition of 5'ppp-RNA in chickens (Gallus gallus), a species lacking RIG-I. Additionally, the authors documented that MDA5's functionality can be compromised by m6A-mediated methylation and degradation of MDA5 mRNA, orchestrated by the METTL3/14-YTHDF2/3 regulatory network in Miiuy croaker during SCRV infection. This impairment compromises the innate antiviral immunity of fish, facilitating SCRV's immune evasion. These findings offer valuable insights into the adaptation and functional diversity of innate antiviral mechanisms in vertebrates.

We extend our sincere appreciation for your professional comments and insightful suggestions on our manuscript, as they have significantly contributed to enhancing its quality.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

(1) The interpretation of Figures 1H and I, along with the captions, seems unclear. Particularly, understanding the meaning of the X-axis in Figure I is challenging. Additionally, the designation of "H2O = 1" on the Y-axis in Figure 1E lacks clarity. It would be helpful if the author could revise and clarify these figures for better comprehension.

We appreciate your reminder and have corrected and clarified these figures and figure legends (lines 768-772). We have replaced the Y-axis of Figure 1I with "Relative mRNA expression" instead of " Relative IFN-1 expression" (Figure 1I). In addition, we have added an explanation of "H2O=1" in the legend of Figure 1E.

(2) The interpretation of Figure 5 in section 2.5 seems incomplete. The author mentioned that both m6A levels and MDA5 expression levels are increased (lines 256-257), prompting questions about the relationship between m6A and MDA5 expression. If higher m6A levels typically lead to MDA5 mRNA instability and lower MDA5 expression, observing both increasing simultaneously appears contradictory. Considering the dynamic changes shown in Figure 5, it would be more appropriate to propose an alteration in both m6A levels and MDA5 expression levels. Given the fluctuating nature of these changes, definitively labeling them as solely "increased" is challenging. Therefore, offering a nuanced interpretation of the results and clarifying this aspect would bolster the study's conclusions.

While changes in m6A modification and the expression of m6A-modified transcripts are biologically relevant, identifying bona fide m6A alterations during viral infection will allow us to understand how m6A modification of cellular mRNA is regulated. As our m6A change analysis pipeline controls for changes in gene expression, these data should represent true changes in m6A modification rather than changes in the expression of m6A-modified transcripts (10.1038/s41598-020-63355-3). Similar studies demonstrated that m6A modification in RIOK3 and CIRBP mRNAs are altered following Flaviviridae infection (10.1016/j.molcel.2019.11.007). The specific calculation method is as follows: relative m6A level for each transcript was calculated as the percent of input in each condition normalized to that of the respective positive control spike-in. Fold change of enrichment was calculated with mock samples normalized to 1. Therefore, the upregulation of MDA5 expression can partially explain the increase in m6A modification on all MDA5 mRNA in cells, but it cannot indicate changes in m6A modification on each mDA5 transcript. We have supplemented the calculation method process in the manuscript and cited relevant literature. I hope to receive your understanding.

In addition, although higher m6A levels often lead to unstable MDA5 mRNA and lower MDA5 expression, SCRV can affect MDA5 expression through multiple pathways. For example, since MDA5 is an interferon-stimulated gene, the infection of SCRV virus can cause strong expression of interferon and indirectly induce high-level expression of MDA5. Therefore, the expression of MDA5 is not contradictory to the simultaneous increase in MDA5 modification (24 h). In order to further enhance our experimental conclusions, we supplemented the dual fluorescence experiment. The results indicate that, the infection of SCRV can inhibit the fluorescence activity of MDA5-exon1 reporter plasmids containing m6A sites but not including the promoter sequence of the MDA5 gene, and this inhibitory effect can be counteracted by cycloleucine (CL, an amino acid analogue that can inhibit m6A modification) (Figure 5G). This further indicates that SCRV can reduce the expression of MDA5 through the m6A pathway.

Finally, in light of the fluctuations in MDA5 expression levels, we have changed the subheadings of Results 2.5 section and provided a more comprehensive and precise elucidation of the experimental outcomes. We are grateful for your valuable feedback.

(3) In the discussion section, it would indeed be advantageous for the author to explore the novelty of this work more comprehensively, moving beyond merely acknowledging the widespread loss of RIG-I and suggesting MDA5 as a compensatory mechanism. Considering the well-established roles of MDA5 and m6A in host-virus interactions, the findings of this study may seem familiar in light of previous research. To enhance the discussion, it would be valuable for the author to delve into the implications of this evolutionary model. For instance, does the compensation or loss of RIG-I impact a species' susceptibility to specific types of viruses? Exploring such questions would provide insight into the broader significance of this compensation model and its potential effects on host-virus interactions, thus adding depth to the study's contribution.

We appreciate the expert advice provided by the referee. In response, we have expanded our discussion in the relevant section, addressing the potential influence of RIG-I deficiency and MDA5 compensation on the antiviral immune system in vertebrates (lines 371-376). Furthermore, we underscore the significance of exploring the impact of SCRV infection on MDA5 m6A modification, considering its compatibility with MDA5 as an ISG gene, in elucidating the host response to viral infection (lines 405-409).

(4) To improve the manuscript, it would be beneficial if the editors could aid the author in refining the language. Many descriptions in the article are overly redundant, and there should be appropriate differentiation between experimental methods and results.

We appreciate the reviewer’s comment. We have carefully revised the manuscript and removed redundant descriptions in the experimental results and methods.

Reviewer #3 (Recommendations For The Authors):

The authors have addressed all of my concerns.

Changes from the earlier version to the revised version

Abstract

Abstract

The incessant arms race between viruses and hosts has led to numerous evolutionary innovations that shape life’s evolution. During this process, the interactions between viral receptors and viruses have garnered significant interest since viral receptors are cell surface proteins exploited by viruses to initiate infection. Our study sheds light on the arms race between the MDA5 receptor and 5’ppp-RNA virus in a lower vertebrate fish, M. miiuy. Firstly, the frequent and independent loss events of RIG-I in vertebrates prompted us to search for alternative immune substitutes, with homology-dependent genetic compensation response (HDGCR) being the main pathway. Our further analysis suggested that MDA5 of M. miiuy and G. gallus, the homolog of RIG-I, can replace RIG-I in recognizing 5’ppp-RNA virus, which may lead to redundancy of RIG-I and loss from the species genome during evolution. Secondly, as an adversarial strategy, 5’ppp-RNA SCRV can utilize the m6A methylation mechanism to degrade MDA5 and weaken its antiviral immune ability, thus promoting its own replication and immune evasion. In summary, our study provides a snapshot into the interaction and coevolution between vertebrate and virus, offering valuable perspectives on the ecological and evolutionary factors that contribute to the diversity of the immune system.

Introduction

The Pattern recognition receptors (PRRs) recognize conserved pathogen-associated molecular pattern (PAMP) motifs, including proteins, lipids and nucleotides, resulting in activation of innate immune responses of host1. Retinoic acid-induced RIG-I-like receptors (RLRs) are a specific group of PRRs belonging to the family of DExD/H-box RNA helicases, and they play a crucial role in the innate antiviral response. There are three members in the RLR family: retinoic acid-inducible gene I (RIG-I), melanoma differentiation-associated gene 5 (MDA5), and laboratory of genetics and physiology 2 (LGP2). RIG-I and MDA5 exhibit similar signaling characteristics and structural homology, including two N-terminal caspase-recruitment domains (CARDs), a DExD/H box RNA helicase domain, and a C-terminal repressor domain (RD)2. However, LGP2 lacks CARD domain2. The specific roles of RIG-I and MDA5 in response to virus stimulation are not redundant: RIG-I detects short blunt-end RNA with a 5’-triphosphate motif (5’ppp-RNA), while MDA5 specifically recognizes long forms of viral dsRNA3–5. In contrast, the involvement of LGP2 in cytosolic RNA sensing remains a topic of debate. Some studies have proposed that LGP2 is crucial for the production of type Ⅰ interferon (IFN-Ⅰ) in response to several RIG-I– and MDA5-dependent viruses6, whereas others have described LGP2 as a negative regulator of RIG-I signaling7. However, surprisingly, despite the important immune functions of RLR members, they have been lost from the genome in many species including mammals, birds, and fish8–10.

5’ppp-RNA is present in a wide variety of RNA viruses, including Flaviviridae, Paramyxoviridae, Coronaviridae, Orthomyxoviridae, and Rhabdoviridae families4. These viruses widely infect all kinds of life and cause various harmfulness and serious diseases, and the infected hosts range from the highest vertebrates to the lowest vertebrate, teleost fish 11, 12. 5’ppp-RNA viruses are expected to be recognized and trigger an antiviral immune response by RIG-I. However, RIG-I is lacking in many vertebrates, so who then recognizes 5’ppp-RNA in these species? When a gene is functionally inactive or lost, other genes or pathways may undergo homology-dependent genetic compensation response (HDGCR) to maintain organismal balance. The pressure from viral infections has led to the evolution of the immune system, and the loss of RIG-I may be accompanied by functional substitution by homologous genes in order to compensate for the loss of RIG-I function. Tupaia belangeri, a mammal lacking RIG-I, has been reported to possess the MDA5 protein, which has undergone positive selection and exhibits the ability to sense Sendai virus (an RNA virus posed as a RIG-I agonist) for inducing type I IFN8. This study provides insights into the compensatory mechanism of RIG-I deficiency in vertebrates, however, considering the dynamic nature of host-virus dynamics, further research is urgently needed to validate and expand these insights.

For millions of years, viruses have coevolved with their hosts, leading to the development of various mechanisms for immune evasion in response to the mutual evolutionary pressure13. Recently, it is reported that N6-methyladenosine (m6A) modification plays a crucial role in enabling viruses to evade the interferon (IFN) response14. m6A methylation is an important epigenetic modification that involves the dynamic and reversible process of RNA methylation, regulated by methylase (writer), demethylase (eraser), and reader proteins that recognize m6A-modified RNA15. N6-methylation deposition on mRNA primarily relies on the enzymatic “writing” complex within the cell, with a key composition featuring methyltransferase-like 3 (METTL3) and METTL14. Conversely, N6-demethylation is chiefly modulated by “eraser” proteins, with fat mass and obesity-associated protein (FTO) and AlkB homolog 5 (ALKBH5) being the identified enzymes for this role 16,17. The functional impact of N6-methylation on RNA characteristics typically remains latent until recognition by “reader” proteins, such as members of the YTHDF protein family, encompassing YTH domain family 1 (YTHDF1), YTHDF2, and YTHDF3. In mammals, YTHDF2 and YTHDF3 are recognized for their involvement in the degradation of m6A-modified RNA targets, as well as their association with mRNA localization and phase separation 18. By utilizing these m6A-related enzymes in the host, viruses can not only use m6A modification to hide themselves and avoid recognition by host receptors but also manipulate host gene expression to enable immune evasion 19–21.

Fish and birds are commonly observed to exhibit “RIG-I deficiency” phenomena, and they are also the most abundant groups of aquatic and terrestrial vertebrates respectively, giving them certain evolutionary advantages. As a result, the research of the immune evolution of them helps us to understand the immune evolution of the entire vertebrates. Meanwhile, fish and birds are major source of high-quality protein for human consumption, viral diseases in aquaculture and livestock industry pose serious obstacles to the industry’s healthy development. Among them, the 5’ppp-RNA virus, as mentioned earlier, is a highly pathogenic group. Within this group, siniperca cheats rhabdovirus (SCRV) is an important pathogen that can cause huge economic losses and serious threats to aquaculture in freshwater and marine fish 22. In addition, vesicular stomatitis virus (VSV) is also a highly infectious and widely prevalent 5’ppp-RNA virus 12. Therefore, in this study, we selected teleost fish miiuy croaker (Miichthys miiuy) and bird chicken (Gallus gallus), two vertebrates lacking RIG-I, as well as two 5’ppp-RNA viruses (SCRV and VSV), to investigate the evolutionary strategies and interactions between 5’ppp-RNA virus and host. The final results demonstrated that the MDA5 of M. miiuy and G. gallus lacking RIG-I both evolved the ability to recognize 5’ppp-RNA. In addition, SCRV virus can in turn utilize the m6A strategy to degrade MDA5 of M. miiuy for immune evasion. In summary, these findings shed light on the functional diversity of innate antiviral activity and unveil a complex arms race between virus replication and the innate immunity of its reservoir hosts in vertebrates. In addition, since the loss of RIG-I in vertebrates is a widespread event, our research also provides a new perspective for vertebrate groups that have lost RIG-I.

Results

A genetic compensation response model for RIG loss

Earlier version: In previous studies, it has been reported that a few vertebrates lost the RIG-I gene during evolution 9, 10, 23. To investigate this further, we utilized the recent advancements in available and comprehensive genomic data of vertebrates from the NCBI and Ensemble websites. By analyzing the genomes of various species, we determined the presence and absence of the RIG-I gene. As depicted in Figure 1A, the absence of the RIG-I gene does not appear to be uncommon in vertebrates, with at least three independent loss events observed in mammals, birds and fish, respectively. These gene-loss events occurred over a long evolutionary timescale. Based on the minimum age of the common ancestor of RIG-I-missing species (http://www.timetree.org/), the oldest recorded losses occurring in species like Erpetoichthys calabaricus, can be traced back to approximately 330 million years ago (MYA). Conversely, the detected loss in Gallus gallus took place less than 42 MYA. To gain a deeper understanding of the evolutionary dynamics of vertebrate RIG-I, we explored the conservation of the genes surrounding RIG-I. We analyzed the genome sequences of Homo sapiens, Callorhinchus milii, Danio rerio, Tupaia belangeri, Gallus gallus, and M. miiuy (Figure 1B). Similar to the slow-evolving Callorhinchus milii, humans possess RIG-I with TOPORS and ACO1 as its upstream and downstream genes, respectively. As previously reported8, 10. Tupaia belangeri, Gallus gallus, and M. miiuy do not have RIG-I. Interestingly, zebrafish has a RIG-I with its upstream gene being TOPORSb, seemingly a result of chromosome fragment duplication. However, one RIG-I gene was lost in the downstream region of TOPORSa, indicating that one of the RIG-I copies has been lost in zebrafish, similar to M. miiuy.

Revised version: In previous studies, it has been reported that a few vertebrates lost the RIG-I gene during evolution 9, 10, 23. To investigate this further, we analyzed the genomes of various species, and determined the presence and absence of the RIG-I gene. As depicted in Figure 1A, the absence of the RIG-I gene does not appear to be uncommon in vertebrates, with at least three independent loss events observed in mammals, birds and fish, respectively. These gene-loss events occurred over a long evolutionary timescale. Based on the minimum age of the common ancestor of RIG-I-missing species (http://www.timetree.org/), the oldest recorded losses occurring in species like Erpetoichthys calabaricus, can be traced back to approximately 330 million years ago (MYA). Conversely, the detected loss in Gallus gallus took place less than 42 MYA. To gain a deeper understanding of the evolutionary dynamics of vertebrate RIG-I, we explored the conservation of the genes surrounding RIG-I. We analyzed the genome sequences of Homo sapiens, Callorhinchus milii, Danio rerio, Tupaia belangeri, Gallus gallus, and M. miiuy (Figure 1B). Similar to the slow-evolving Callorhinchus milii, humans possess RIG-I with TOPORS and ACO1 as its upstream and downstream genes, respectively. As previously reported8, 10. Tupaia belangeri, Gallus gallus, and M. miiuy do not have RIG-I. Interestingly, zebrafish has a RIG-I with its upstream gene being TOPORSb, seemingly a result of chromosome fragment duplication. However, one RIG-I gene was lost in the downstream region of TOPORSa, indicating that one of the RIG-I copies has been lost in zebrafish, similar to M. miiuy.

A genetic compensation response model for RIG loss.

Earlier version: (A) Loss of RIG-I in vertebrates. Each branch tip represents one species. Red lines show lineages where loss of RIG-I. The red circle indicates the occurrence of an independent loss event. Branch lengths scale in millions of years (MYA). (B) Comparative analysis of gene synteny of RIG-I in vertebrate genomes. (C and D) SCRV-RNA and VSV-RNA were extracted from SCRV and VSV virus and dephosphorylated them with Calf Intestinal Alkaline Phosphatase (CIAP). Then, MKC cells were transfected with SCRV-RNA and SCRV-RNA+CIAP, and DF-1 cells were transfected with VSV-RNA and VSV-RNA+CIAP. IRF3 dimerization was analyzed by native gel electrophoresis. (E) MKC cells were transfected with different concentrations of SCRV-RNA and SCRV-RNA+CIAP, then the expression of IFN-1 was detected by qRT-PCR. (F) Predicted protein structures of MDA5 and LGP2 of M. miiuy and RIG-I of H. sapiens. (G) The expression of MDA5 and LGP2 plasmids of M. miiuy were detected by Western Blot. (H and I) MKC cells were co-transfected with MDA5 plasmids and LGP2 plasmids (H) or si-LGP2 (I) for 24 h and then treated with SCRV (MOI = 5) for 24 h. The expression of IFN-1 was determined by qPCR. (J) A genetic compensation response model for RIG loss. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01 versus the controls.

Revised version: (A) Loss of RIG-I in vertebrates. Each branch tip represents one species. Red lines show lineages where loss of RIG-I. The red circle indicates the occurrence of an independent loss event. Branch lengths scale in millions of years (MYA). (B) Comparative analysis of gene synteny of RIG-I in vertebrate genomes. (C and D) SCRV-RNA and VSV-RNA were extracted from SCRV and VSV virus and dephosphorylated them with Calf Intestinal Alkaline Phosphatase (CIAP). Then, MKC cells were transfected with SCRV-RNA and SCRV-RNA+CIAP, and DF-1 cells were transfected with VSV-RNA and VSV-RNA+CIAP. IRF3 dimerization was analyzed by native gel electrophoresis. (E) MKC cells were transfected with different concentrations of SCRV-RNA and SCRV-RNA+CIAP, then the expression of IFN-1 was detected by qRT-PCR. As a control, the expression level of IFN-1 in MKC cells transfected with the same volume of water was set to “1”. (F) Predicted protein structures of MDA5 and LGP2 of M. miiuy and RIG-I of H. sapiens. (G) The expression of MDA5 and LGP2 plasmids of M. miiuy were detected by Western Blot. (H) MKC cells were co-transfected with MDA5 plasmids and LGP2 plasmids for 24 h and then treated with SCRV (MOI = 5) for 24 h. The expression of IFN-1 was determined by qPCR. (I) MKC cells were co-transfected with MDA5 plasmids and si-LGP2 for 24 h and then treated with SCRV (MOI = 5) for 24 h. The expression of LGP2 and IFN-1 was determined by qPCR. (J) A genetic compensation response model for RIG loss. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01 versus the controls.

Earlier version: Given the important role of RIG-I in recognizing 5’ppp RNA, we selected two different models for experiments to detect the presence of other potential receptors for immune compensation in RIG-I deficient species. Specifically, we extracted virus RNA from SCRV and VSV viruses and obtained their dephosphorylated RNA derivatives (SCRV-RNA-CIAP and VSV-RNA-CIAP) through Calf Intestinal Alkaline Phosphatase (CIAP) treatment. Next, these RNAs were transfected into MKC cells of M. miiuy and DF-1 cells of G. gallus respectively to detect the dimerization form of IRF3. Results showed that virus RNA are more potent than dephosphorylated virus RNA in activating IRF3 (Figure 1C and 1D). And further qPCR experiments showed that the induction ability of IFN-1 by SCRV-RNA is significantly higher than that of SCRV-RNA-CIAP. These results indicate the presence of viral receptors in M. miiuy and G. gallus that can recognize 5’ppp-RNA. Next, we compared the structural domains of human RIG-I and M. miiuy MDA5 (mmiMDA5) and LGP2 (mmiLGP2) proteins, and observed that mmiLGP2 lacks the CARD domain (Figure 1F). We next constructed the mmiLGP2 plasmid and synthesized siRNA targeting mmiLGP2, and co-transfected them with MDA5 plasmids into MKC cells. The results showed that mmiLGP2 can actively regulate mmiMDA5’s immune response to the SCRV virus (Figure 1H and 1I), which further indicates that it participates in antiviral immunity as an intermediate regulatory factor of MDA5 rather than a downstream signal transmitter7. Therefore, these results prompted us to investigate whether MDA5, a structurally similar homolog to RIG-I, can substitute for RIG-I in recognizing 5’ppp-RNA in vertebrates without RIG-I (Figure 1J).

Revised version: Given the important role of RIG-I in recognizing 5’ppp RNA, we selected two different models for experiments to detect the presence of other potential receptors for immune compensation in RIG-I deficient species. Specifically, we extracted virus RNA from SCRV and VSV viruses and obtained their dephosphorylated RNA derivatives (SCRV-RNA-CIAP and VSV-RNA-CIAP) through Calf Intestinal Alkaline Phosphatase (CIAP) treatment. Next, these RNAs were transfected into MKC cells of M. miiuy and DF-1 cells of G. gallus respectively to detect the dimerization form of IRF3. Results showed that virus RNA are more potent than dephosphorylated virus RNA in activating IRF3 (Figure 1C and 1D). And further qPCR experiments showed that the induction ability of IFN-1 by SCRV-RNA is significantly higher than that of SCRV-RNA-CIAP (Figure 1E). These results indicate the presence of viral receptors in M. miiuy and G. gallus that can recognize 5’ppp-RNA. Next, we compared the structural domains of human RIG-I and M. miiuy MDA5 (mmiMDA5) and LGP2 (mmiLGP2) proteins, and observed that mmiLGP2 lacks the CARD domain (Figure 1F). We next co-transfected mmiLGP2 plasmids or si-mmiLGP2 with MDA5 plasmids into MKC cells. The results showed that mmiLGP2 can actively regulate mmiMDA5’s immune response to the SCRV virus (Figure 1G-1I), which further indicates that it participates in antiviral immunity as an intermediate regulatory factor of MDA5 rather than a downstream signal transmitter7. Therefore, these results prompted us to investigate whether MDA5, a structurally similar homolog to RIG-I, can substitute for RIG-I in recognizing 5’ppp-RNA in vertebrates without RIG-I (Figure 1J).

MDA5 promotes host antiviral innate immunity

Earlier version: To further investigate the alternative immune function of MDA5, we first investigated the M. miiuy MDA5-mediated signaling pathway in response to SCRV virus and double-stranded mimetic poly(I:C)-HMW infection. Two siRNA molecules (si-MDA5-1 and si-MDA5-2) were designed to assess the function of MDA5. si-MDA5-1 demonstrated higher inhibitory efficiency compared to si-MDA5-2 (Figure 2A). Therefore, si-MDA5-1 (referred to as si-MDA5) was selected for subsequent experiments. Subsequently, we constructed the MDA5 expression plasmid and confirmed its protein overexpression efficiency and cytoplasmic localization (Figure 2B and 2C). Considering that mammalian MDA5 activates NF-κB and IRF3/IRF7 to induce IFNs, we investigated whether M. miiuy MDA5 could impact these pathways. Dual-luciferase reporter assays demonstrated that knockdown of MDA5 significantly inhibits NF-κB, IRF3, IRF7, and IFN-1 reporter genes (Figure 2D). In addition, we investigated the role of MDA5 in regulating the expression of IFN-stimulated genes (ISGs), which are crucial effectors. As shown in Figure 2E-2G, we observed that knockdown or overexpression of MDA5 can significantly inhibit or promote the expression levels of IFN-1, Mx1, ISG15, and Viperin upon SCRV infection or poly(I:C)-HMW stimulation. To visually represent the biological significance of MDA5 in SCRV-induced host cells, we utilized the virus plaque and qPCR methods to examine whether it could influence viral replication. As demonstrated in Figure 2H-2K, knockdown of MDA5 led to a significant increase in viral plaque formation and promoted SCRV replication in SCRV-infected cells, while overexpression had the opposite effect. Next, we explored that how MDA5 conducts signal transduction, and results showed that MDA5 is capable of binding not only to MAVS but also to STING (Figure 2L and 2M). Under the infection of SCRV virus, co-transfection of STING plasmids or si-STING with MDA5 plasmids can significantly increase or decrease the induction of IFN-1, indicating the important role of STING in MDA5-mediated immune transmission in response to RNA virus (Figure 2N and 2O). Finally, through the determination of IRF3 dimer and IFN induction, our results showed that the induction of antiviral innate immunity by SCRV-RNA depends on MDA5 (Figure 2P and 2Q). Overall, these findings indicated that MDA5 plays a positive role in regulating the antiviral responses of M. miiuy.

Revised version: To further investigate the alternative immune function of MDA5, we first investigated the M. miiuy MDA5-mediated signaling pathway in response to SCRV virus and double-stranded mimetic poly(I:C)-HMW infection. Two siRNA molecules (si-MDA5-1 and si-MDA5-2) were designed to assess the function of MDA5. si-MDA5-1 demonstrated higher inhibitory efficiency compared to si-MDA5-2 (Figure 2A). Therefore, si-MDA5-1 (referred to as si-MDA5) was selected for subsequent experiments. Subsequently, we constructed the MDA5 expression plasmid and confirmed its protein overexpression efficiency and cytoplasmic localization (Figure 2B and 2C). Considering that mammalian MDA5 activates NF-κB and IRF3/IRF7 to induce IFNs, we investigated whether M. miiuy MDA5 could impact these pathways. Dual-luciferase reporter assays demonstrated that knockdown of MDA5 significantly inhibits NF-κB, IRF3, IRF7, and IFN-1 reporter genes (Figure 2D). In addition, we investigated the role of MDA5 in regulating the expression of IFN-stimulated genes (ISGs), which are crucial effectors. As shown in Figure 2E-2G, we observed that knockdown or overexpression of MDA5 can significantly inhibit or promote the expression levels of IFN-1, Mx1, ISG15, and Viperin upon SCRV infection or poly(I:C)-HMW stimulation. To visually represent the biological significance of MDA5 in SCRV-induced host cells, we utilized the virus plaque and qPCR methods to examine whether it could influence viral replication. As demonstrated in Figure 2H-2K, knockdown of MDA5 led to a significant increase in viral plaque formation and promoted SCRV replication in SCRV-infected cells, while overexpression had the opposite effect. Next, we explored that how MDA5 conducts signal transduction, and results showed that MDA5 is capable of binding not only to MAVS but also to STING (Figure 2L and 2M). In our previous laboratory investigations, we have substantiated the induction effect of STING on IFN under SCRV infection or poly(I:C) stimulation 24. Under the infection of SCRV virus, co-transfection of STING plasmids or si-STING with MDA5 plasmids can significantly increase or decrease the induction of IFN-1, indicating the important role of STING in MDA5-mediated immune transmission in response to RNA virus (Figure 2N and 2O). Finally, through the determination of IRF3 dimer and IFN induction, our results showed that the induction of antiviral innate immunity by SCRV-RNA depends on MDA5 (Figure 2P and 2Q). Overall, these findings indicated that MDA5 plays a positive role in regulating the antiviral responses of M. miiuy.

MDA5 promotes host antiviral innate immunity.

Earlier version: (A) Silence efficiency of si-MDA5 measured by qRT-PCR and Western blotting. Two siRNAs of MDA5 were transfected into MPC cells for 48 h respectively. (B) MKC cells were transfected with MDA5 plasmids for 48 h, then the overexpression efficiency was detected through MDA5 endogenous antibodies. (C) Subcellular localization of MDA5 in MKCs by immunofluorescence. (D) MDA5 knockdown suppresses NF-κB, IRF3, IRF7, and IFN-1 signaling. MPC cells were transfected with pRL-TK Renilla luciferase plasmid, luciferase reporter genes, together with MDA5 expression plasmid. At 48 h post-transaction, the luciferase activity was measured and normalized to renilla luciferase activity. (E) Knockdown of MDA5 attenuates the expression of antiviral genes. MPC cells were transfected with si-Ctrl or si-MDA5 for 24 h and then treated with SCRV (MOI =5) or poly(I:C)-HMW for 24 h. The expression of IFN-1, Mx1, ISG15, and Viperin were determined by qPCR. (F and G) Overexpression of MDA5 promotes the expression of antiviral genes. MKC cells were transfected with vector or MDA5 expression plasmids for 24 h and then treated with SCRV (MOI = 5) or poly (I:C)-HMW for 24 h. The expression of MDA5 (F), and antiviral genes, including IFN-1, Mx1, ISG15, and Viperin (G) were determined by qPCR. (H and J) MPC cells transfected with si-Ctrl or si-MDA5 and MKC cells transfected with pcDNA3.1 vector or MDA5 plasmid 24 h and then treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (I and K) MDA5 inhibits SCRV replication. MPC cells were transfected with si-Ctrl or si-MDA5 and MKC cells were transfected with pcDNA3.1 vector or MDA5 expression plasmid for 24 h, then infected with SCRV (MOI = 5) for 12, 24, or 48 h. The qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. (L and M) MDA5 immunoprecipitates with MAVS (L) and STING (M). EPC cells (∼ 1×107) were co-transfected with MAVS-Myc or STING-Myc and MDA5-Flag expression plasmids for 24 h, followed by immunoprecipitation (IP) with anti-Myc. (N and O) MKC cells were transfected with MDA5 plasmids and STING plasmids (N) or si-STING (O) for 24 h and then treated with SCRV (MOI = 5) for 24 h. The expression of IFN-1 was determined by qPCR. (P and Q) MKC cells were transfected with MDA5 plasmids and si-STING for 24 h and then transfected with SCRV-RNA for 24 h, then the IRF3 dimerization and IFN-1 levels were analyzed by native gel electrophoresis (P) and qPCR (Q), respectively. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01; *, p < 0.05 versus the controls.

Revised version: (A) Silence efficiency of si-MDA5 measured by qRT-PCR and Western blotting. Two siRNAs of MDA5 were transfected into MPC cells for 48 h respectively. (B) MKC cells were transfected with MDA5 plasmids for 48 h, then the overexpression efficiency was detected through MDA5 endogenous antibodies. (C) Subcellular localization of MDA5 in MKCs by immunofluorescence. (D) MDA5 knockdown suppresses NF-κB, IRF3, IRF7, and IFN-1 signaling. MPC cells were transfected with pRL-TK Renilla luciferase plasmid, luciferase reporter genes, together with MDA5 expression plasmid. At 48 h post-transaction, the luciferase activity was measured and normalized to renilla luciferase activity. (E) Knockdown of MDA5 attenuates the expression of antiviral genes. MPC cells were transfected with si-Ctrl or si-MDA5 for 24 h and then treated with SCRV (MOI =5) or poly(I:C)-HMW for 24 h. The expression of IFN-1, Mx1, ISG15, and Viperin were determined by qPCR. (F and G) Overexpression of MDA5 promotes the expression of antiviral genes. MKC cells were transfected with vector or MDA5 expression plasmids for 24 h and then treated with SCRV (MOI = 5) or poly (I:C)-HMW for 24 h. The expression of MDA5 (F), and antiviral genes, including IFN-1, Mx1, ISG15, and Viperin (G) were determined by qPCR. (H and J) MPC cells transfected with si-Ctrl or si-MDA5 and MKC cells transfected with pcDNA3.1 vector or MDA5 plasmid 24 h and then treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (I and K) MDA5 inhibits SCRV replication. MPC cells were transfected with si-Ctrl or si-MDA5 and MKC cells were transfected with pcDNA3.1 vector or MDA5 expression plasmid for 24 h, then infected with SCRV (MOI = 5) for 12, 24, or 48 h. The qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. (L and M) MDA5 immunoprecipitates with MAVS (L) and STING (M). EPC cells (∼ 1×107) were co-transfected with MAVS-Myc or STING-Myc and MDA5-Flag expression plasmids for 24 h, followed by immunoprecipitation (IP) with anti-Myc. (N and O) MKC cells were transfected with MDA5 plasmids and STING plasmids (N) or si-STING (O) for 24 h and then treated with SCRV (MOI = 5) for 24 h. The expression of IFN-1 was determined by qPCR. (P and Q) MKC cells were transfected with si-MDA5 for 24 h and then transfected with SCRV-RNA for 24 h, then the IRF3 dimerization and IFN-1 levels were analyzed by native gel electrophoresis (P) and qPCR (Q), respectively. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01; *, p < 0.05 versus the controls.

RD domain is required for MDA5 to recognize SCRV RNA

Earlier version: We investigated whether M. miiuy MDA5, in the absence of RIG-I, can directly bind to the RNA of SCRV. To address this question, we precipitated the MDA5 protein from SCRV-infected MKC cells that were transfected with MDA5-Flag and pcDNA3.1-Flag. The bound RNA was then amplified using SCRV-specific primers or β-actin primers as a control. Our results showed that MDA5 interacted with SCRV RNA, while β-actin mRNA did not associate with MDA5 (Figure 3A). The C-terminal regulatory domain (RD) of human RIG-I binds to viral RNA in a triphosphate-dependent manner, to determine if the RD domain of M. miiuy MDA5 has similar functions, we generated truncated mutations called MDA5-ΔRD and Western blot analysis confirmed the expression of MDA5 and MDA5-ΔRD plasmids (Figure 3B). We subsequently noted that MDA5-ΔRD exhibited an inability to interact with SCRV RNA, which was supported by the loss of its immunological capability in inhibiting SCRV viral replication (Figure 3C-3E). We also examined the effects of MDA5 and its mutant on antiviral factors in both poly(I:C)-HMW-stimulated and SCRV-infected MKC cells. Results showed that only the overexpression of MDA5, not MDA5-ΔRD, significantly increased the expression of antiviral genes (Figure 3F and 3G). In summary, these findings provided evidence that MDA5 may recognize 5’ppp-RNA through its RD domain.

Revised version: We investigated whether M. miiuy MDA5, in the absence of RIG-I, can directly bind to the RNA of SCRV. To address this question, we precipitated the MDA5 protein from SCRV-infected MKC cells that were transfected with MDA5-Flag and pcDNA3.1-Flag. The bound RNA was then amplified using SCRV-specific primers or β-actin primers as a control. Our results showed that MDA5 interacted with SCRV RNA, while β-actin mRNA did not associate with MDA5 (Figure 3A). The C-terminal regulatory domain (RD) of human RIG-I binds to viral RNA in a triphosphate-dependent manner, to determine if the RD domain of M. miiuy MDA5 has similar functions, we generated truncated mutations called MDA5-ΔRD and Western blot analysis confirmed the expression of MDA5 and MDA5-ΔRD plasmids (Figure 3B). We subsequently noted that MDA5-ΔRD exhibited an inability to interact with SCRV RNA, which was supported by the loss of its immunological capability in inhibiting SCRV viral replication (Figure 3C-3E). We also examined the effects of MDA5 and MDA5-ΔRD on antiviral factors in both uninfected, SCRV-infected, and poly(I:C)-HMW-stimulated MKC cells. In the absence of infection, overexpression of both MDA5 and MDA5-ΔRD stimulated the expression of antiviral genes (Figure 3F-3I). However, when cells were infected or stimulated with SCRV or poly(I:C)-HMW, only the overexpression of MDA5, not MDA5-ΔRD, significantly increased the expression of antiviral genes (Figure 3F-3I). In summary, these findings provided evidence that MDA5 may recognize SCRV RNA through its RD domain.

RD domain is required for MDA5 to recognize SCRV RNA.

Earlier version: (A) The association of MDA5 proteins with SCRV. (B) Schematics and expression of MDA5 and truncated MDA5 (MDA5-ΔRD). (C) The association of MDA5 and MDA5-ΔRD proteins with SCRV. (D) MKC cells transfected with pcDNA3.1 vector, MDA5, or MDA5-ΔRD plasmids for 24 h and then treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (E) MKC cells were transfected with pcDNA3.1 vector, MDA5 or MDA5-ΔRD expression plasmids for 24 h, then infected with SCRV for 24 h, then the qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. (F and G) Expression antiviral genes in poly(I:C)-HMW-stimulated (F) or SCRV-infected (G) MKC cells transfected with pcDNA3.1 vector, MDA5, or MDA5-ΔRD expression plasmids. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01; *, p < 0.05 versus the controls.

Revised version: (A) The association of MDA5 proteins with SCRV. (B) Schematics and expression of MDA5 and truncated MDA5 (MDA5-△RD). (C) The association of MDA5 and MDA5-△RD proteins with SCRV. (D) MKC cells transfected with pcDNA3.1 vector, MDA5, or MDA5-△RD plasmids for 24 h and then treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (E) MKC cells were transfected with pcDNA3.1 vector, MDA5 or MDA5-△RD expression plasmids for 24 h, then infected with SCRV for 24 h, then the qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. (F-I) Expression of IFN-1 (F), Mx1 (G), ISG15 (H), and Viperin (I) in Mock-infected, SCRV-infected, and poly(I:C)-HMW-stimulated MKC cells transfected with pcDNA3.1 vector, MDA5, or MDA5-△RD expression plasmids. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01 versus the controls.

MDA5 recognizes 5’ppp-RNA in M. miiuy and G. gallus

Earlier version: To provide specific evidence of MDA5 recognition of 5’ppp-RNA in M. miiuy, we synthesized 67 nt 5’ppp-SCRV from the SCRV virus and through in vitro transcription using T7 polymerase. As a positive control capable of binding to the mmiMDA5 receptor, 112 bp dsRNA from the MDA5 gene sequence was also generated. Gel analysis and nuclease sensitivity confirmed the generation of RNA products of the expected size (Figure 4A). The RNA pull-down experiment showed that 5’ppp-SCRV can be recognized by mmiMDA5, similar to 112 bp blunt end dsRNA (Figure 4B). Next, we transfected wild-type 5’ppp-SCRV and dephosphorylated mutant SCRV RNAs (5’-OH-SCRV and 5’pppGG-SCRV) into MKC cells to detect the dimer of IRF3, and found that the activation ability of 5’ppp-RNA on IRF3 depends on its triphosphate structure (Figure 4C). Furthermore, biotin-labeled 5’ppp-RNA captured mmiMDA5, while the biotin-labeled mutant types (5’-OH-SCRV and 5’pppGG-SCRV) did not (Figure 4D). Additionally, we purified recombinant MDA5 and MDA5-ΔRD of M. miiuy to conduct electrophoresis mobility shift assay (EMSA) (Figure 4E). When recombinant mmiMDA5 was subjected to EMSA using a biotin-labeled 5’ppp-SCRV probe, a clear complex was detected (Figure 4F). This complex formation was inhibited by adding an excess of cold 5’ppp-SCRV but not 5’-OH-SCRV and 5’pppGG-SCRV, demonstrating the direct interaction between mmiMDA5 and 5’ppp-SCRV (Figure 4F). However, 5’ppp-SCRV cannot bind to the MDA5-ΔRD (Figure 4G). Next, we investigated whether G. gallus MDA5 (ggaMDA5) shares the ability to recognize 5’ppp-RNA seen in fish. We synthesized 67 nt 5’ppp-VSV from the VSV virus and found that it is similar to 5’ppp-SCRV, binding to ggaMDA5 and activating IRF3 in a triphosphate-dependent manner (4H-4L). Overall, our findings demonstrated that MDA5 detection of RNA in two RIG-I deficient vertebrates (M. miiuy and G. gallus) is guided by the RNA 5’-triphosphate end and is disrupted when RNA is capped at the 5’-triphosphate end or dephosphorylated.

Revised version: To provide specific evidence of MDA5 recognition of 5’ppp-RNA in M. miiuy, we synthesized 67 nt 5’ppp-SCRV from the SCRV virus and through in vitro transcription using T7 polymerase. As a positive control capable of binding to the mmiMDA5 receptor, 112 bp dsRNA from the MDA5 gene sequence was also generated. Gel analysis and nuclease sensitivity confirmed the generation of RNA products of the expected size (Figure 4A). The RNA pull-down experiment showed that 5’ppp-SCRV can be recognized by mmiMDA5, similar to 112 bp blunt end dsRNA (Figure 4B). Next, we transfected wild-type 5’ppp-SCRV and dephosphorylated mutant SCRV RNAs (5’-OH-SCRV and 5’pppGG-SCRV) into MKC cells to detect the dimer of IRF3, and found that the activation ability of 5’ppp-RNA on IRF3 depends on its triphosphate structure (Figure 4C). in addition, knockdown of mmiMDA5 blocked the induction of IRF3 by 5’ppp-SCRV (Figure 4D). Furthermore, biotin-labeled 5’ppp-RNA captured mmiMDA5, while the biotin-labeled mutant types (5’-OH-SCRV and 5’pppGG-SCRV) did not (Figure 4E). Additionally, we purified recombinant MDA5 and MDA5-ΔRD of M. miiuy to conduct electrophoresis mobility shift assay (EMSA) (Figure 4F). When recombinant mmiMDA5 was subjected to EMSA using a biotin-labeled 5’ppp-SCRV probe, a clear complex was detected (Figure 4G). This complex formation was inhibited by adding an excess of cold 5’ppp-SCRV but not 5’-OH-SCRV and 5’pppGG-SCRV, demonstrating the direct interaction between mmiMDA5 and 5’ppp-SCRV (Figure 4G). However, 5’ppp-SCRV cannot bind to the MDA5-ΔRD (Figure 4H). Next, we investigated whether G. gallus MDA5 (ggaMDA5) shares the ability to recognize 5’ppp-RNA seen in fish. We synthesized 67 nt 5’ppp-VSV from the VSV virus and found that it is similar to 5’ppp-SCRV, binding to ggaMDA5 and activating IRF3 in a triphosphate-dependent manner (4I-4M). Overall, our findings demonstrated that MDA5 detection of RNA in two RIG-I deficient vertebrates (M. miiuy and G. gallus) is guided by the RNA 5’-triphosphate end and is disrupted when RNA is capped at the 5’-triphosphate end or dephosphorylated.

MDA5 recognizes 5’ppp-RNA in M. miiuy and G. gallus.

Earlier version: (A) Schematic representation of 5’ppp-SCRV and dsRNA. The product of in vitro transcription runs as a single product degraded by RNase I. (B) Pull-down of biotinylated or non-biotinylated dsRNA and 5’ppp-SCRV by mmiMDA5. MKC cells were transfected with mmiMDA5-Flag plasmid, and the input and immunoprecipitated MDA5 proteins were analyzed by Western blot. (C) MKC cells were transfected with 5’ppp-SCRV, 5’OH-SCRV (5’ppp-SCRV dephosphorylated by CIAP), or 5’pppGG-SCRV (5’ppp-SCRV capped by m7G cap analog), then IRF3 dimerization was analyzed by native gel electrophoresis. (D) The cytoplasmic fraction of MKC cells transfected with Flag-tagged mmiMDA5 was incubated with biotinylated 5’ppp-SCRV, 5’OH-SCRV, or 5’pppGG-SCRV. RNA-protein complexes were pulled down using streptavidin affinity beads. Input and pull-down samples were analyzed by SDS-PAGE and immunoblotting using anti-Flag antibody. (E) Purity of recombinant mmiMDA5 and mmiMDA5-ΔRD were determined by SDS-PAGE and immunoblotting using anti-His antibody. (F) EMSA of 5’ppp-SCRV with mmiMDA5. For binding competition, indicated non-biotinylated RNAs (50-fold molar excess over the probe) were included. (G) EMSA of 5’ppp-SCRV with mmiMDA5 or mmiMDA5-ΔRD. (H) Schematic representation of 5’ppp-VSV. The product of in vitro transcription runs as a single product degraded by RNase I. (I) Pull-down of dsRNA and 5’ppp-VSV by ggaMDA5. DF-1 cells were transfected with ggaMDA5-Flag plasmids, and the input and immunoprecipitated MDA5 proteins were analyzed by Western blot. (J) DF-1 cells were transfected with 5’ppp-VSV, 5’OH-VSV (5’ppp-VSV dephosphorylated by CIAP), or 5’pppGG-VSV (5’ppp-VSV capped by m7G cap analog), then IRF3 dimerization was analyzed by native gel electrophoresis. (K)The cytoplasmic fraction of DF-1 cells transfected with Flag-tagged ggaMDA5 was incubated with biotinylated 5’ppp-VSA, 5’OH-VSV, or 5’pppGG-RNA. RNA-protein complexes were pulled down using streptavidin affinity beads. Input and pull-down samples were analyzed by SDS-PAGE and immunoblotting using anti-Flag antibody. (L) EMSA of 5’ppp-RNA with ggaMDA5. For binding competition, indicated non-biotinylated RNAs (50-fold molar excess over the probe) were included.

Revised version: (A) Schematic representation of 5’ppp-SCRV and dsRNA. The product of in vitro transcription runs as a single product degraded by RNase I. (B) Pull-down of biotinylated or non-biotinylated dsRNA and 5’ppp-SCRV by mmiMDA5. MKC cells were transfected with mmiMDA5-Flag plasmid, and the input and immunoprecipitated MDA5 proteins were analyzed by Western blot. (C) MKC cells were transfected with 5’ppp-SCRV, 5’OH-SCRV (5’ppp-SCRV dephosphorylated by CIAP), or 5’pppGG-SCRV (5’ppp-SCRV capped by m7G cap analog), then IRF3 dimerization was analyzed by native gel electrophoresis. (D) WT and si-MDA5 MKC cells were transfected with 5’ppp-SCRV and poly(I:C)-HMW, then IRF3 dimerization was analyzed by native gel electrophoresis. (E) The cytoplasmic fraction of MKC cells transfected with Flag-tagged mmiMDA5 was incubated with biotinylated 5’ppp-SCRV, 5’OH-SCRV, or 5’pppGG-SCRV. RNA-protein complexes were pulled down using streptavidin affinity beads. Input and pull-down samples were analyzed by SDS-PAGE and immunoblotting using anti-Flag antibody. (F) Purity of recombinant mmiMDA5 and mmiMDA5-△RD were determined by SDS-PAGE and immunoblotting using anti-His antibody. (G) EMSA of 5’ppp-SCRV with mmiMDA5. For binding competition, indicated non-biotinylated RNAs (50-fold molar excess over the probe) were included. The asterisk marks the binding bands between mmiMDA5 protein and different RNA probes. (H) EMSA of 5’ppp-SCRV with mmiMDA5 or mmiMDA5-△RD. The asterisk marks the binding bands between mmiMDA5 or mmiMDA5-△RD proteins and 5’ppp-SCRV RNA probes. (I) Schematic representation of 5’ppp-VSV. The product of in vitro transcription runs as a single product degraded by RNase I. (J) Pull-down of dsRNA and 5’ppp-VSV by ggaMDA5. DF-1 cells were transfected with ggaMDA5-Flag plasmids, and the input and immunoprecipitated MDA5 proteins were analyzed by Western blot. (K) DF-1 cells were transfected with 5’ppp-VSV, 5’OH-VSV (5’ppp-VSV dephosphorylated by CIAP), or 5’pppGG-VSV (5’ppp-VSV capped by m7G cap analog), then IRF3 dimerization was analyzed by native gel electrophoresis. (L) The cytoplasmic fraction of DF-1 cells transfected with Flag-tagged ggaMDA5 was incubated with biotinylated 5’ppp-VSA, 5’OH-VSV, or 5’pppGG-RNA. RNA-protein complexes were pulled down using streptavidin affinity beads. Input and pull-down samples were analyzed by SDS-PAGE and immunoblotting using anti-Flag antibody. (M) EMSA of 5’ppp-RNA with ggaMDA5. For binding competition, indicated non-biotinylated RNAs (50-fold molar excess over the probe) were included. The asterisk marks the binding bands between ggaMDA5 protein and different RNA probes.

Added in the revised version: Altered m6A modification and expression of MDA5 upon SCRV infection

Added in the revised version: As stated in the previous results, MDA5 has evolved to possess immune recognition abilities that take over the role of RIG-I, significantly enhancing the host’s ability to resist viral infections. However, evolution is a continuous battle between hosts and viruses, viruses often evolve complex escape mechanisms to evade immune surveillance and destruction, in which the m6A modification mechanism is an important way. To investigate the impact of SCRV infection on the overall methylation dynamics of fish hosts, we utilized methylated RNA immunoprecipitation sequencing (MeRIP-seq, GenBank accession number: PRJNA819945) to measure changes in m6A modification of host transcripts during SCRV infection (Figure 5A). From the results shown in Figure 5B, we identified 3060 differentially m6A methylated peaks (p-value < 0.05) and 2082 differentially expressed genes (p-value < 0.05) between the non-infected and SCRV-infected group, indicating that infection significantly alters the m6A modification of specific host transcripts. It is worth noting that hyper-methylated m6A peaks were also observed in the MDA5 mRNA transcripts during SCRV infection (fold change >1.5, p < 0.05), and the expression level of MDA5 was higher in the infection group compared to the control group (p < 0.05). Further analysis of the m6A-seq data showed that MDA5 contained m6A sites in the exon1 region which were markedly higher in the infection group than that in the control group (Untreated VS SCRV: p = 0.01) (Figure 5C). Next, we tested the concrete m6A modification sites of MDA5 mRNA. We used the SRAMP database to predict potential m6A sites in MDA5-exon1 25. Four potential sites were predicted and two specific primers were designed (Figure 5D), then m6A quantitative real-time PCR and semiquantitative PCR were performed. In line with the m6A-seq data, the results further confirmed that m6A sites exactly exist in MDA5 exon1 (Figure 5E). After being infected with SCRV or poly(I:C)-HMW, the methylation level of MDA5 increased in MKC cells, further confirming the sequencing data (Figure 5F). We further inserted the MDA5-exon1 into pmirGLO, and results showed that the infection of SCRV can reduce the luciferase activity of pmirGLO-MDA5-exon1, and the use of cycloleucine (CL, an amino acid analogue that can inhibit m6A modification) hindered the degradation of MDA5-exon1 by SCRV, indicating that SCRV can regulate MDA5 expression through a methylation mechanism (Figure 5G). MDA5 is classified as an interferon-stimulated gene (ISG). Consistently, in spleen and kidney tissues infected with SCRV, the expression level of MDA5 demonstrates fluctuation. (Figure 5H). We further explored the mRNA and protein levels of MDA5 in MKC cells infected with SCRV and poly(I:C)-HMW and found that their expression all first significantly increased and then returned to the normal level (Figure 5I and 5J). Overall, our results suggested that the m6A modification of MDA5 is increased, and the expression level of MDA5 shows fluctuating changes upon SCRV infection.

Increased m6A modification and expression of MDA5 upon SCRV infection.

Removed in the revised version: As stated in the previous results, MDA5 has evolved to possess immune recognition abilities that take over the role of RIG-I, significantly enhancing the host’s ability to resist viral infections. However, evolution is a continuous battle between hosts and viruses, viruses often evolve complex escape mechanisms to evade immune surveillance and destruction, in which the m6A modification mechanism is an important way. To investigate the impact of SCRV infection on the overall methylation dynamics of fish hosts, we utilized methylated RNA immunoprecipitation sequencing (MeRIP-seq, GenBank accession number: PRJNA819945) to measure changes in m6A modification of host transcripts during SCRV infection (Figure 5A). From the results shown in Figure 5B, we identified 3060 differentially m6A methylated peaks (p-value < 0.05) and 2082 differentially expressed genes (p-value < 0.05) between the non-infected and SCRV-infected group, indicating that infection significantly alters the m6A modification of specific host transcripts. It is worth noting that hyper-methylated m6A peaks were also observed in the MDA5 mRNA transcripts during SCRV infection (fold change >1.5, p < 0.05), and the expression level of MDA5 was higher in the infection group compared to the control group (p < 0.05). Further analysis of the m6A-seq data showed that MDA5 contained m6A sites in the exon1 region which were markedly higher in the infection group than that in the control group (Untreated VS SCRV: p = 0.01) (Figure 5C). Next, we tested the concrete m6A modification sites of MDA5 mRNA. We used the SRAMP database to predict potential m6A sites in MDA5-exon1 24. Four potential sites were predicted and two specific primers were designed (Figure 5D), then m6A quantitative real-time PCR and semiquantitative PCR were performed. In line with the m6A-seq data, the results further confirmed that m6A sites exactly exist in MDA5 exon1 (Figure 5E). After being infected with SCRV or poly(I:C)-HMW, the methylation level of MDA5 increased in MKC cells, further confirming the sequencing data (Figure 5F). Next, MDA5 was found to be upregulated in spleen and kidney tissues infected with SCRV (Figure 5G). We further explored the mRNA and protein levels of MDA5 in MKC cells infected with SCRV and poly(I:C)-HMW and found that their expression all first significantly increased and then returned to the normal level (Figure 5H and 5I). Overall, our results suggested that both the m6A modification and expression levels of MDA5 are increased upon SCRV infection.

Removed in the revised version: Increased m6A modification and expression of MDA5 upon SCRV infection.

Earlier version: (A) Schematic of the MeRIP-seq protocol used to identify differential m6A methylation following infection of M. miiuy spleen tissues with SCRV. (B) Circos plots showing differentially m6A-methylated peaks identified from normal and SCRV-infected spleen tissues of M. miiuy. The green and red lines on the innermost ring represented the m6A peaks identified from normal and SCRV-infected spleen tissues of M. miiuy respectively. The orange lines represent the log2 fold change values for each differentially m6A peak. The blue spots represent the fold change of gene expression. Chromosomes were shown on the outermost ring, the ruler on which represented the physical distance is in millions of bases (Mb). (C) The m6A abundance in MDA5 mRNA transcripts detected by MeRIP-seq, the m6A peak of MDA5 is circled in the green box. (D) The exon1 sequence of MDA5 was submitted to the SRAMP website, and then the predicted m6A site was displayed. Four predicted methylation sites located on the exon1 of MDA5 were marked by red box, then we designed two m6A specific primers (MDA5-m6A-1 and MDA5-m6A-2). (E) m6A abundance on MDA5 exon1 detected by MeRIP-qPCR and emiquantitative PCR in MKC cells. (F) MeRIP-qPCR analysis of relative m6A level of MDA5 in Mock, SCRV-infected, and poly(I:C)-HMW-stimulated MKC cells. (G) mRNA levels of MDA5 in spleen and kidney tissues measured by qRT-PCR at indicated time after SCRV infection (MOI = 5). (H and I) mRNA and protein levels of MDA5 in MKC cells measured by qRT-PCR and Western blotting at indicated time after SCRV infection (MOI = 5) (H) and poly(I:C)-HMW stimulation (I). All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01 versus the controls.

Revised version: (A) Schematic of the MeRIP-seq protocol used to identify differential m6A methylation following infection of M. miiuy spleen tissues with SCRV. (B) Circos plots showing differentially m6A-methylated peaks identified from normal and SCRV-infected spleen tissues of M. miiuy. The green and red lines on the innermost ring represented the m6A peaks identified from normal and SCRV-infected spleen tissues of M. miiuy respectively. The orange lines represent the log2 fold change values for each differentially m6A peak. The blue spots represent the fold change of gene expression. Chromosomes were shown on the outermost ring, the ruler on which represented the physical distance is in millions of bases (Mb). (C) The m6A abundance in MDA5 mRNA transcripts detected by MeRIP-seq, the m6A peak of MDA5 is circled in the green box. (D) The exon1 sequence of MDA5 was submitted to the SRAMP website, and then the predicted m6A site was displayed. Four predicted methylation sites located on the exon1 of MDA5 were marked by red box, then we designed two m6A specific primers (MDA5-m6A-1 and MDA5-m6A-2). (E) m6A abundance on MDA5 exon1 detected by MeRIP-qPCR and emiquantitative PCR in MKC cells. (F) MeRIP-qPCR analysis of relative m6A level of MDA5 in Mock, SCRV-infected, and poly(I:C)-HMW-stimulated MKC cells. (G) Relative luciferase activity of pmirGLO-MDA5-exon1 firefly luciferase reporters in Mock and SCRV-infected (MOI = 5) MKC cells that treated with DMSO or Cycloleucine (CL, 20 mM). (H) mRNA levels of MDA5 in spleen and kidney tissues measured by qRT-PCR at indicated time after SCRV infection (MOI = 5). (I and J) mRNA and protein levels of MDA5 in MKC cells measured by qRT-PCR and Western blotting at indicated time after SCRV infection (MOI = 5) (I) and poly(I:C)-HMW stimulation (J). All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01 versus the controls.

m6A-modification weakens MDA5 mRNA stability and antiviral ability

Earlier version: SCRV infection can increase the methylation level of MDA5, and this mechanism is expected to have a regulatory effect on MDA5. To investigate this, the regulatory effects of METTL3 and METTL14, the main components of the methyltransferase complex, on MDA5 were explored. Three siRNAs targeting METTL3 and METTL14 were designed, and si-METTL3-2 and si-METTL14-3 showed higher knockdown efficiency (Figure 6A). METTL3 and METTL14-overexpressing plasmids were also constructed, and the interaction between them confirmed their cooperative role in m6A modification (Figure 6B). Subsequently, the global m6A levels were measured in cells after knockdown or overexpression of METTL3 and METTL14. Consistent with expectations, the knockdown of METTL3 and METTL14 significantly decreased the global m6A level, while their overexpression significantly increased the overall methylation level (Figure 6C). To investigate their role in the process of virus infection, we used the virus plaque and qPCR methods to explore whether it could affect SCRV replication. As shown in Figure 6D and 6E, overexpression of METTL3 or METTL14 significantly enhances viral plaque formation and promotes SCRV replication in SCRV-infected cells. And as m6A writers, METTL3 and METTL14 can directly regulate the methylation level of MDA5 and induce MDA5 degradation by reducing its mRNA stability (Figure 6F-6I). We further inserted the MDA5-exon1 into pmirGLO and mVenus-C1 vector, and results showed that si-METTL3&14 can reduce the luciferase activity of pmirGLO-MDA5-exon1, and METTL3&14 could suppress the levels of GFP, further indicating the presence of methylation sites in MDA5-exon1 (Figure 6-figure supplement 1A and 1D). To investigate the potential sites on MDA5-exon1, we constructed mutant MDA5-exon1 reporter plasmids. For the mutant form of MDA5-exon1 reporter genes, four predicted adenosine bases in m6A consensus sequences (DRACH) were replaced by uracil, thus canceling m6A decoration (Figure 5D and Figure 6-figure supplement 1B). As expected, our results showed that the luciferase activity of wild-type pmirGLO-MDA5-exon1 reporter, but not of mutant pmirGLO-MDA5-exon1 reporter, was markedly increased upon METTL3&14 knockdown (Figure 6-figure supplement 1C). Furthermore, we found that the use of cycloleucine (CL, an amino acid analogue that can inhibit m6A modification) hindered the degradation of MDA5 by METTL3 and METTL14, indicating that these proteins regulate MDA5 expression through a methylation mechanism (Figure 6J and 6K). Additionally, MDA5 expression was degraded by METTL3&14 in both SCRV-infected and uninfected groups (Figure 6L). To investigate the impact of this mechanism on viral infection, we co-transfected METTL3&14 and MDA5 plasmids and found that METTL3&14 partially suppressed the antiviral immune benefit of MDA5 (Figure 6M-6O). Taken together, our data demonstrate that METTL3&14-mediated m6A methylation of MDA5 inhibits its mRNA stability and expression levels, thus suppressing the antiviral capabilities of MDA5.

Revised version: SCRV infection can increase the methylation level of MDA5, and this mechanism is expected to have a regulatory effect on MDA5. To investigate this, the regulatory effects of METTL3 and METTL14, the main components of the methyltransferase complex, on MDA5 were explored. Three siRNAs targeting METTL3 and METTL14 were designed, and si-METTL3-2 and si-METTL14-3 showed higher knockdown efficiency (Figure 6A). METTL3 and METTL14-overexpressing plasmids were also constructed, and the interaction between them confirmed their cooperative role in m6A modification (Figure 6B). Subsequently, the global m6A levels were measured in cells after knockdown or overexpression of METTL3 and METTL14. Consistent with expectations, the knockdown of METTL3 and METTL14 significantly decreased the global m6A level, while their overexpression significantly increased the overall methylation level (Figure 6C). To investigate their role in the process of virus infection, we used the virus plaque and qPCR methods to explore whether it could affect SCRV replication. As shown in Figure 6D and 6E, overexpression of METTL3 or METTL14 significantly enhances viral plaque formation and promotes SCRV replication in SCRV-infected cells. And as m6A writers, METTL3 and METTL14 can directly regulate the methylation level of MDA5 and induce MDA5 degradation by reducing its mRNA stability (Figure 6F-6I). We further inserted the MDA5-exon1 into pmirGLO and mVenus-C1 vector, and results showed that si-METTL3&14 can reduce the luciferase activity of pmirGLO-MDA5-exon1, and METTL3&14 could suppress the levels of GFP, further indicating the presence of methylation sites in MDA5-exon1 (Figure 6-figure supplement 1A and 1D). To investigate the potential m6A sites on MDA5-exon1, we constructed mutant MDA5-exon1 reporter plasmids (Figure 5D and Figure 6-figure supplement 1B). As expected, our results showed that the luciferase activity of wild-type pmirGLO-MDA5-exon1 reporter, but not of mutant pmirGLO-MDA5-exon1 reporter, was markedly increased upon METTL3&14 knockdown (Figure 6-figure supplement 1C). Furthermore, we found that the use of cycloleucine hindered the degradation of MDA5 by METTL3 and METTL14, indicating that these proteins regulate MDA5 expression through a methylation mechanism (Figure 6J and 6K). Additionally, MDA5 expression was degraded by METTL3&14 in both SCRV-infected and uninfected groups (Figure 6L). To investigate the impact of this mechanism on viral infection, we co-transfected METTL3&14 and MDA5 plasmids and found that METTL3&14 partially suppressed the antiviral immune benefit of MDA5 (Figure 6M-6O). Taken together, our data demonstrate that METTL3&14-mediated m6A methylation of MDA5 inhibits its mRNA stability and expression levels, thus suppressing the antiviral capabilities of MDA5.

m6A-modification weakens MDA5 mRNA stability and antiviral ability.

(A) Silence efficiency of si-METTL3 and si-METTL14 measured. Three siRNAs of METTL3 and METTL14 were transfected into MPC cells for 48 h respectively. (B) METTL3 immunoprecipitates with METTL14. EPC cells (1×107) were co-transfected with METTL3-Myc and METTL14-Flag expression plasmids for 24 h, followed by immunoprecipitation (IP) with anti-Myc. (C) METTL3 and METTL14 significantly increased the m6A content. MPCs were transfected with si-Ctrl or si-METTL3 or si-METTL14 and MKCs were transfected with vector or METTL3 or METTL14 plasmids for 48 h, then the m6A level was measured by colorimetry. (D) METTL3 and METTL14 overexpressed MKC cells seeded in 48-well plates overnight were treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (E) MKC cells were transfected with pcDNA3.1 vector and METTL3 or METTL14 expression plasmid for 24 h, then infected with SCRV (MOI = 5) for 24 h. The qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. (F) The m6A level alteration of MDA5 upon METTL3 or METTL14 knockdown or overexpression was examined by MeRIP-qPCR. MPC cells were transfected with si-Ctrl or si-METT3 and MKC cells were transfected with vector or METTL3 or METTL14 plasmids for 48 h. (G and H) MPC cells were transfected with si-Ctrl, si-METTL3 or METTL14 and MKC cells were transfected with vector, METTL3, or METTL14 plasmids for 48 h, the expression of MDA5 was detected by qRT-PCR and Western blotting. (I) MPC cells were transfected with si-Ctrl, si-METTL3, or si-METTL14, and MKC cells were transfected with vector, METTL3, or METTL14 plasmids, then 5 µg/ml actinomycin D was added to the cells for 0 h, 2 h, and 4 h. The half-life of MDA5 was analyzed by qRT-PCR. (J) MKC cells were transfected with vector or METTL3&14 plasmids for 24 h and then treated with Cycloleucine (CL) for 24 h in a final concentration of 20 mM. The expression of MDA5 was detected by qRT-PCR and Western blotting. (K) MKC cells were transfected with vector or METTL3&14 plasmids for 24 h and then treated with CL for 24 h at 20 mM, then 5 µg/ml actinomycin D was added to the cells for 0 h, 2 h, and 4 h. The half-life of MDA5 was analyzed by qRT-PCR. (L) MKC cells were transfected with vector or METTL3&14 plasmids for 24 h and stimulated with SCRV (MOI = 5) for 24 h, then the expression of MDA5 was detected by qRT-PCR and Western blotting. (M) MKC cells seeded in 48-well plates overnight were transfected with MDA5 or MDA5+METTL3&14 plasmids for 48 h, then the expression of MDA5 was detected by qRT-PCR. (N) MKC cells seeded in 48-well plates overnight were transfected with MDA5 or MDA5+METTL3&14 plasmids were treated with SCRV at the dose indicated for 48 h. Then, cell monolayers were fixed with 4% paraformaldehyde and stained with 1% crystal violet. (O) MKC cells were transfected with MDA5 or MDA5 + METTL3&14 plasmids for 24 h, then infected with SCRV (MOI = 5) for 24 h. The qPCR analysis was conducted for SCRV-M and SCRV-G RNA levels. All data presented as the means ± SE from at least three independent triplicated experiments. **, p < 0.01; *, p < 0.05 versus the controls.

Detailed m6A regulatory mechanism of MDA5

The m6A mechanism is jointly regulated by writers, erasers, and readers. The m6A modification of mRNA recruits m6A-binding proteins (readers) which have the function of recognizing m6A sites and accelerating mRNA decay. In this study, we hypothesized that YTHDF readers (YTHDF1, YTHDF2, and YTHDF3) may recognize m6A-modified MDA5 and destabilize it. We examined the binding ability of YTHDF proteins to MDA5 mRNA and found that YTHDF2 and YTHDF3 exhibited stronger binding compared to YTHDF1 (Figure 7A). We then designed two siRNAs targeting YTHDF1, YTHDF2, and YTHDF3 respectively, and selected si-YTHDF1-1, si-YTHDF2-si-2, si-YTHDF3-2 with better knockdown effects for subsequent experiments (Figure 7B). Results showed that the knockdown or overexpression of YTHDF2 and YTHDF3 led to the increase or decrease of the expression levels of MDA5, whereas YTHDF1 did not appear to influence MDA5 expression (Figure 7C, 7D, Figure 7-figure supplement 1A and 1C). In addition, overexpression of YTHDF2&3 can increase the luciferase activity of wild-type pmirGLO-MDA5-exon1 reporter, but not of mutant pmirGLO-MDA5-exon1 reporter (Figure 7-figure supplement 1B). We then found that si-METTL3&14 can counteract the negative impact of YTHDF2&3 on MDA5 mRNA, protein, and stability (Figure 7E and 7F). Since the m6A mechanism is reversible, we then studied the regulation of MDA5 by two major erasers. Results showed that FTO, not ALKBH5, can upregulate MDA5 expression (Figure 7G). Moreover, the overexpression of FTO can reverse the inhibitory effect of YTHDF2&3 on MDA5 (Figure 7H and 7I). Taken together, these data suggested that through METTL3&14-m6A-YTHDF2&3 regulatory network, m6A can mediate the degradation of MDA5, and this process can be countered by FTO (Figure 7J).

Detailed m6A regulatory mechanism of MDA5.

(A) The binding relationship between YTHDF1, YTHDF2, or YTHDF3 and MDA5 mRNA was validated using a RIP assay. MKC cells were transfected with YTHDF1-Flag, YTHDF2-Flag, YTHDF3-Flag, or pcDNA3.1-Flag for 48 h. (B) SiRNA silencing effect test of YTHDF1, YTHDF2, and YTHDF3. MPC cells were transfected with si-YTHDF1, si-YTHDF2, or si-YTHDF3 for 48 h. (C and D) MPC cells were transfected with si-YTHDF1, si-YTHDF2, si-YTHDF3 or si-Ctrl (C), and MKC cells were transfected with YTHDF1, YTHDF2, YTHDF3 or vector for 48 h (D), then the expression of MDA5 was detected by qRT-PCR and Western blotting. (E) Relative mRNA and protein levels of MDA5 in MKC cells after co-transfected with YTHDF2&3 plasmids and si-METTL3&14 by qPCR and western blot assays. (F) MKC cells were transfected with YTHDF2&3 plasmids and si-METTL3&14, then 5 µg/ml actinomycin D was added to the cells for 0 h, 2 h, and 4 h. The half-life of MDA5 was analyzed by qRT-PCR. (G) MKC cells were transfected with vector, FTO, ALKBH5 plasmids for 48 h, then the expression of MDA5 was detected by qRT-PCR and Western blotting. (H) Relative mRNA and protein levels of MDA5 in MKC cells after co-transfected with YTHDF2&3 plasmids and FTO plasmids by qPCR and western blotting. (I) MKC cells were transfected with YTHDF2&3 plasmids and FTO plasmids, then 5 µg/ml actinomycin D was added to the cells for 0 h, 2 h, and 4 h. The half-life of MDA5 was analyzed by qRT-PCR. (J) Schematic diagram of arms race between MDA5 and 5’ppp-RNA virus in M. miiuy.

Discussion

In this study, we have identified multiple independent loss events of RIG-I throughout vertebrate evolution. Specifically, we have observed at least three independent loss events in mammals, birds, and fish, respectively. Regarding the loss of RIG-I in avian species, previous studies have provided more comprehensive and accurate information than ours, indicating that birds have experienced more than 16 instances of RIG-I loss10. However, it should be noted that we have not found the loss of RIG-I in reptiles and amphibians (not displayed). But this does not imply that there has been no loss event among them, but rather, there is insufficient information regarding the genomes of reptiles and amphibians to draw the conclusions. As the abundance of genomic data continues to grow, there will be an increasing amount of evidence available to verify the presence or absence of RIG-I in these taxonomic groups. Additionally, the presence of RIG-I in current species does not guarantee that RIG-I loss has not occurred previously. For instance, our research has indicated that zebrafish may have experienced RIG-I loss. However, prior to the loss event, the zebrafish RIG-I has undergone repetitive events. These findings collectively highlight that the loss of RIG-I is not a random occurrence, but rather bears significant biological and evolutionary implications that deserve further investigation.

Earlier version: The loss of RIG-I might be expected to have a great impact on immune response, and the evolutionary loss of it raises the question of a possible compensation of its function by other gene products. RIG-I deficiency instances in mammals were scarcer compared to those in fish and birds. Previous research has revealed that in T. belangeri, concurrent with the loss of RIG-I, both MDA5 and LGP2 have undergone notable positive selection, and MDA5 or MDA5/LGP2 could sense Sendai virus (an RNA virus posed as a RIG-I agonist) for inducing type I IFN. This proposed an immune substitution model for RIG-I deficiency in mammals, though direct evidence of MDA5/LGP2 binding to 5’ppp-RNA was lacking. In this study, we characterized the consequences resulting from the loss of RIG-I in the M. miiuy and G. gallus (two representative species of vertebrates lacking RIG-I) and provided direct functional evidence for the alternative immune recognition function of MDA5. SCRV and VSV belongs to the class of rhabdovirus, which is a negative-strand virus with a non-segmented genome and initiate both replication and transcription de novo leading to 5′-triphosphate RNA in the cytosol, was expected to trigger an IFN response without the need for replication and presumed dsRNA formation4. Then, we synthesized 5’ppp-RNA probes from the SCRV and VSV genome sequences and purified recombinant MDA5 of M. miiuy and G. gallus. Results demonstrated that MDA5 detect the 5’ppp-RNA in M. miiuy and G. gallus and is disrupted when RNA is capped at the 5’-triphosphate end or dephosphorylated. Furthermore, the recognition of RNA by MDA5 depends on its RD domain, which not only indicates that the RD domain of MDA5 has evolved functions similar to RD domain of RIG-I, but also excludes the possibility of MDA5 binding to 5’ppp-RNA through intermediate ligands. Furthermore, in vertebrates that possess RIG-I, STING exclusively binds to RIG-I rather than MDA5. However, similar to the interaction observed between MDA5 and STING in tree shrews and chickens that lacking RIG-I8, 25, the MDA5 of M. miiuy can also interact with STING protein, indicating that in vertebrates lacking RIG-I, MDA5 may utilize STING to facilitate signal transduction in the antiviral response. Finally, by linking the immune replacement ability evolved by MDA5 with the common RIG-I loss event in vertebrates, we propose a straightforward explanation for this evolutionary event: In the lengthy process of vertebrate evolution, MDA5 gradually acquired new functionalities to effectively detect viruses such as 5’ppp-RNAs, which were initially recognized by RIG-I. As a result, there was functional redundancy of RIG-I, which eventually led to its gradual disappearance from the vertebrate genome. This hypothesis is supported by two viewpoints: first, MDA5 and RIG-I can both recognize overlapping viruses26; secondly, in fish without RIG-I deficiency, MDA5 has stronger antiviral ability compared to RIG-I27.

Revised version: The evolutionary loss of RIG-I raises the question of a possible compensation of its function by other gene products. RIG-I deficiency instances in mammals were scarcer compared to those in fish and birds. Previous research has revealed that in T. belangeri, concurrent with the loss of RIG-I, both MDA5 and LGP2 have undergone notable positive selection, and MDA5 or MDA5/LGP2 could sense Sendai virus (an RNA virus posed as a RIG-I agonist) for inducing type I IFN. This proposed an immune substitution model for RIG-I deficiency in mammals, though direct evidence of MDA5/LGP2 binding to 5’ppp-RNA was lacking. In this study, we characterized the consequences resulting from the loss of RIG-I in the M. miiuy and G. gallus (two representative species of vertebrates lacking RIG-I) and provided direct functional evidence for the alternative immune recognition function of MDA5. SCRV and VSV belongs to the class of rhabdovirus, which is a negative-strand virus with a non-segmented genome and initiate both replication and transcription de novo leading to 5′-triphosphate RNA in the cytosol, was expected to trigger an IFN response without the need for replication and presumed dsRNA formation4. Then, we synthesized 5’ppp-RNA probes from the SCRV and VSV genome sequences and purified recombinant MDA5 of M. miiuy and G. gallus. Results demonstrated that MDA5 detect the 5’ppp-RNA in M. miiuy and G. gallus and is disrupted when RNA is capped at the 5’-triphosphate end or dephosphorylated. Furthermore, the recognition of RNA by MDA5 depends on its RD domain, which not only indicates that the RD domain of MDA5 has evolved functions similar to RD domain of RIG-I, but also excludes the possibility of MDA5 binding to 5’ppp-RNA through intermediate ligands. Furthermore, in vertebrates that possess RIG-I, STING exclusively binds to RIG-I rather than MDA5. However, similar to the interaction observed between MDA5 and STING in tree shrews and chickens that lacking RIG-I8, 26, the MDA5 of M. miiuy can also interact with STING protein, indicating that in vertebrates lacking RIG-I, MDA5 may utilize STING to facilitate signal transduction in the antiviral response. Next, by linking the immune replacement ability evolved by MDA5 with the common RIG-I loss event in vertebrates, we propose a straightforward explanation for this evolutionary event: In the lengthy process of vertebrate evolution, MDA5 gradually acquired new functionalities to effectively detect viruses such as 5’ppp-RNAs, which were initially recognized by RIG-I. As a result, there was functional redundancy of RIG-I, which eventually led to its gradual disappearance from the vertebrate genome. This hypothesis is supported by two viewpoints: first, MDA5 and RIG-I can both recognize overlapping viruses27; secondly, in fish without RIG-I deficiency, MDA5 has stronger antiviral ability compared to RIG-I28. Finally, although MDA5 has a compensatory effect on RIG-I, the absence of RIG-I may still have a significant impact on the immune system of vertebrates. Ducks are often resistant to influenza viruses capable of killing chicken, and the loss of RIG-I in chickens provides a plausible explanation for their increased susceptibility to influenza viruses compared with ducks 29. Furthermore, SCRV virus, the virus model used in this study, mainly affects the Siniperca chuatsi, which is a fish species that also lacks RIG-I like M. miiuy 30.

Earlier version: Viruses are constantly co-evolving with hosts. Our available data suggested that the vertebrate has evolved MDA5 with alternative functions to recognize and resist 5’ppp-RNA virus in the absence of RIG-I. As a response, virus has also evolved to use the m6A mechanism to evade the host immune response. Previous study indicated that m6A modification plays a crucial role in maintaining immune homeostasis in the body28. During the infection process, the virus may disrupt the m6A level and disrupt the immune balance, thereby facilitating its own replication and invasion. Viruses may utilize the m6A mechanism in two ways to evade immunity. On the one hand, many studies suggested a role for m6A in shielding RNA species from detection by PRRs. There are numerous viruses recognized in RIG-I dependent manner, but m6A modification increases the possibility of immune evasion. For example, the loss of m6A site or removing m6A from human metapneumovirus may weaken its immune evasion ability, prevent it from evading detection by RIG-I, and result in higher levels of IFN release20. Moreover, it has been reported that the RNA containing m6A modifications binds to RIG-I poorly, and it failed to trigger the RIG-I conformational changes associated with downstream immune signals29. On the other hand, m6A-mediated viral strategies for inhibiting IFN induction pathways appear to exist. For instance, following infection by human cytomegalovirus (HCMV), the IFN-β transcript was modified by m6A, and this methylation decreased its half-life, suggesting that HCMV can exploit this control of IFN-β expression to facilitate its replication by increasing m6A modification of the IFN-β transcript and thus decreasing its production30. In another study, the m6A reader YTHDF3 could restrain the type I IFN response and ISG expression by involving enhancement of FOXO3 translation in a m6A-independent mechanism, which acts as a negative regulator of immunity 19. Moreover, previous study reports that IFN-α and IFN-β mRNA encoding cytokines that promote IFN response are modified by m6A, and the loss of METTL3 and YTHDF2 could increase the IFN expression and ISG activity 31. In summary, viruses can utilize m6A to evade the immune response by shielding the recognition of RIG-I or reducing the host’s immune reaction, thus counteracting the host’s defense mechanisms. In our previous study, we discovered that SCRV infection can not only increase overall cellular methylation levels but also upregulate the expression of METTL3, thereby suggesting its potential to regulate host gene expression through the m6A mechanism32. In this study, we found that SCRV virus infection alters m6A modification of specific cellular transcripts, including MDA5. MDA5 plays a crucial role as a receptor in recognizing the SCRV virus and controlling IFN activation. The aberrant m6A levels of MDA5 subsequently affect the expression of antiviral genes, thereby compromising its ability to mount an effective antiviral immune response. It is plausible that SCRV also impacts the production of antiviral genes, such as IFN-1 or other cytokines, by manipulating m6A modification on the transcripts of cytokines or molecules involved in their production. This interesting aspect warrants further investigation in future research.

Revised version: Viruses are constantly co-evolving with hosts. Our available data suggested that the vertebrate has evolved MDA5 with alternative functions to recognize and resist 5’ppp-RNA virus in the absence of RIG-I. As a response, virus has also evolved to use the m6A mechanism to evade the host immune response. Previous study indicated that m6A modification plays a crucial role in maintaining immune homeostasis in the body31. During the infection process, the virus may disrupt the m6A level and disrupt the immune balance, thereby facilitating its own replication and invasion. Viruses may utilize the m6A mechanism in two ways to evade immunity. On the one hand, many studies suggested a role for m6A in shielding RNA species from detection by PRRs. There are numerous viruses recognized in RIG-I dependent manner, but m6A modification increases the possibility of immune evasion. For example, the loss of m6A site or removing m6A from human metapneumovirus may weaken its immune evasion ability, prevent it from evading detection by RIG-I, and result in higher levels of IFN release20. Moreover, it has been reported that the RNA containing m6A modifications binds to RIG-I poorly, and it failed to trigger the RIG-I conformational changes associated with downstream immune signals32. On the other hand, m6A-mediated viral strategies for inhibiting IFN induction pathways appear to exist. For instance, following infection by human cytomegalovirus (HCMV), the IFN-β transcript was modified by m6A, and this methylation decreased its half-life, suggesting that HCMV can exploit this control of IFN-β expression to facilitate its replication by increasing m6A modification of the IFN-β transcript and thus decreasing its production33. In another study, the m6A reader YTHDF3 could restrain the type I IFN response and ISG expression by involving enhancement of FOXO3 translation in a m6A-independent mechanism, which acts as a negative regulator of immunity 19. Moreover, previous study reports that IFN-α and IFN-β mRNA encoding cytokines that promote IFN response are modified by m6A, and the loss of METTL3 and YTHDF2 could increase the IFN expression and ISG activity 34. In summary, viruses can utilize m6A to evade the immune response by shielding the recognition of RIG-I or reducing the host’s immune reaction, thus counteracting the host’s defense mechanisms. In our previous study, we discovered that SCRV infection can not only increase overall cellular methylation levels but also upregulate the expression of METTL3, thereby suggesting its potential to regulate host gene expression through the m6A mechanism35. While changes in m6A modification and the expression of m6A-modified transcripts are biologically relevant, identifying bona fide m6A alterations during viral infection will allow us to understand how m6A modification of cellular mRNA is regulated. As our m6A change analysis pipeline controls for changes in gene expression, these data should represent true changes in m6A modification rather than changes in the expression of m6A-modified transcripts. In this study, we found that SCRV virus infection alters m6A modification of specific cellular transcripts, including MDA5. MDA5 plays a crucial role as a receptor in recognizing the SCRV virus and controlling IFN activation. The aberrant m6A levels of MDA5 subsequently affect the expression of antiviral genes, thereby compromising its ability to mount an effective antiviral immune response. It is plausible that SCRV also impacts the production of antiviral genes, such as IFN-1 or other cytokines, by manipulating m6A modification on the transcripts of cytokines or molecules involved in their production. This interesting aspect warrants further investigation in future research.

In this study, we demonstrated a co-evolutionary arms race between 5’ppp-RNA virus and virus receptors. Specifically, we conducted an evolutionary analysis and provided functional evidence to confirm that MDA5 of M. miiuy and G. gallus has acquired an additional function of sensing 5’ppp-RNA, thus compensating for the loss of RIG-I. Additionally, we found that SCRV infection can regulate the m6A level of MDA5 in the M. miiuy, leading to its degradation and subsequently affecting the immune response for immune evasion. Fish and birds, as the most abundant aquatic and terrestrial vertebrates respectively, are also the two lineages with the most frequent occurrence of RIG-I deficiency, the investigation of host-pathogen interactions of them offers valuable insights into the ecological and evolutionary factors that contribute to the diversity of the immune system.

Materials and methods

Ethics statement

All animal experimental procedures were performed in accordance with the National Institutes of Health’s Guide for the Care and Use of Laboratory Animals, and the experimental protocols were approved by the Research Ethics Committee of Shanghai Ocean University (No. SHOU-DW-2018-047).

Fish and challenge

Earlier version: M. miiuy (∼50 g) was obtained from Zhoushan Fisheries Research Institute, Zhejiang Province, China. Fish was acclimated in aerated seawater tanks at 25°C for six weeks before experiments. The experimental procedure for SCRV infection was performed as described33. Briefly, In the SCRV injection group, fish was challenged with 200 μl SCRV at a multiplicity of infection (MOI) of 5 through intraperitoneal. As a comparison, 200 μl of physiological saline was used to challenge the individuals. Afterwards, fishes were respectively sacrificed at different time point and the kidney and spleen tissues were collected for RNA extraction.

Revised version: M. miiuy (∼50 g) was obtained from Zhoushan Fisheries Research Institute, Zhejiang Province, China. Fish was acclimated in aerated seawater tanks at 25℃ for six weeks before experiments. The experimental procedure for SCRV infection was performed as described36. Briefly, In the SCRV injection group, fish was challenged with 200 μl SCRV at a multiplicity of infection (MOI) of 5 through intraperitoneal. As a comparison, 200 μl of physiological saline was used to challenge the individuals. Afterwards, fishes were respectively sacrificed at different time point and the kidney and spleen tissues were collected for RNA extraction.

Earlier version: Cell culture and treatment

Revised version: Cell culture

Earlier version: M. miiuy kidney cells (MKCs) and spleen cells (MPCs) are cell lines derived from the kidney and spleen tissues of M. miiuy 34, with passages ranging from 20 to 40 times. MKC and MPC cells were cultured in L-15 medium (HyClone) supplemented with 15% fetal bovine serum (FBS; Gibco), 100 U/ml penicillin, and 100 μg/ml streptomycin at 26°C. Fish EPC cells (Epithelioma papulosum cyprini cells) were maintained in medium 199 (Invitrogen) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 26°C in 5% CO2. HEK293T cells were cultured in DMEM (HyClone) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 37°C in 5% CO2. Chicken embryo fibroblast cells (DF-1) were cultured in DMEM (HyClone) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 39°C in 5% CO2. For stimulation experiments, cells were challenged with SCRV at a multiplicity of infection (MOI) of 5 and harvested at different times for RNA extraction.

Revised version: M. miiuy kidney cells (MKCs) and spleen cells (MPCs) are cell lines derived from the kidney and spleen tissues of M. miiuy 37, with passages ranging from 20 to 40 times. MKC and MPC cells were cultured in L-15 medium (HyClone) supplemented with 15% fetal bovine serum (FBS; Gibco), 100 U/ml penicillin, and 100 μg/ml streptomycin at 26℃. Fish EPC cells (Epithelioma papulosum cyprini cells) were maintained in medium 199 (Invitrogen) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 26℃ in 5% CO2. HEK293T cells were cultured in DMEM (HyClone) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 37℃ in 5% CO2. Chicken embryo fibroblast cells (DF-1) were cultured in DMEM (HyClone) supplemented with 10% FBS, 100 U/ml penicillin, and 100 mg/ml streptomycin at 39℃ in 5% CO2.

Oligonucleotides

Earlier version: The sequence of the 62 nt 5’ppp-RNA was derived from the 5’ and 3’UTRs of the VSV genome as previously described37, and the sequence of 112 bp dsRNA was taken from the first 100 bp of the MDA5 gene of miiuy croaker flanked by 5’-gggaga and tctccc-3’. 67 nt 5’ppp-SCRV or 5’ppp-VSV and 112 bp dsRNA were synthesized by in vitro transcription using the Ribo RNAmax-T7 kit (RiboBio). Biotin-labeled RNAs were transcribed in vitro using a Ribo RNAmax-T7 Biotin Labeling Transcription Kit (RiboBio). For generating duplexes, RNA oligonucleotides were mixed in hybridization buffer (20 mM Tris-HCl [pH 8.0], 1.5 mM MgCl2, and 1.5 mM DTT), boiled for 1 min, and incubated at 37°C for 1 h. For removal of 5’-triphosphates, 20 mg of RNA was treated with 20 U Calf Intestinal Alkaline Phosphatase (CIAP, Invitrogen) for 2 h at 37°C in the presence of RNase inhibitor (Beyotime) and extracted with phenolchloroform. RNA molecules with a terminal 5’ m7G cap were synthesized by the incorporation of m7G cap analogue (Yeason) in the transcription reactions. RNA was analysed on a denaturing 17% polyacrylamide, 7 M urea gel following digestion with 50 ng/μl of RNase A (Beyotime) or 100 mU/μl of DNase I (Beyotime) for 30 min.

Revised version: The sequences of the 67 nt 5’ppp-SCRV or 5’ppp-VSV were derived from the 5’ and 3’UTRs of the VSV or SCRV genome as previously described38, and the sequence of 112 bp dsRNA was taken from the first 100 bp of the MDA5 gene of miiuy croaker flanked by 5’-gggaga and tctccc-3’. These RNAs were synthesized by in vitro transcription using the Ribo RNAmax-T7 kit (RiboBio). Biotin-labeled RNAs were transcribed in vitro using a Ribo RNAmax-T7 Biotin Labeling Transcription Kit (RiboBio). For generating duplexes, RNA oligonucleotides were mixed in hybridization buffer (20 mM Tris-HCl [pH 8.0], 1.5 mM MgCl2, and 1.5 mM DTT), boiled for 1 min, and incubated at 37℃ for 1 h. For removal of 5’-triphosphates, 20 mg of RNA was treated with 20 U Calf Intestinal Alkaline Phosphatase (CIAP, Invitrogen) for 2 h at 37℃ in the presence of RNase inhibitor (Beyotime) and extracted with phenolchloroform. RNA molecules with a terminal 5’ m7G cap were synthesized by the incorporation of m7G cap analogue (Yeason) in the transcription reactions. RNA was analysed on a denaturing 17% polyacrylamide, 7 M urea gel following digestion with 50 ng/μl of RNase A (Beyotime) or 100 mU/μl of DNase I (Beyotime) for 30 min.

Plasmids construction

To construct Flag or Myc-tagged MDA5 (GeneBank: PP179381.1), MAVS, STING, METTL3, METTL14, FTO, ALKBH5, YTHDF1, YTHDF2, and YTHDF3 expression plasmids, their CDS sequences of the M. miiuy were amplified by PCR and then ligated into a pcDNA3.1 vector (Invitrogen), respectively. Then, MDA5-△RD were generated by PCR on the basis of MDA5 plasmid. To construct MDA5-exon1 reporter vector, the exon1 region of M. miiuy MDA5 gene was amplified using PCR with the gene-specific primers and ligated into a pmir-GLO luciferase reporter vector (Promega). Meanwhile, M. miiuy MDA5-exon1 were inserted into the mVenus-C1 (Invitrogen), which included the sequence of enhanced GFP. For mutant reporter vectors (MDA5-exon1-mut), adenosine (A) in the predicted m6A motif was replaced by an uracil (U) by using the Mut Express II Fast Mutagenesis Kit V2 (Vazyme) with specific primers). To construct Flag-tagged MDA5 expression plasmid of G. gallus, the CDS sequence of MDA5 was amplified by PCR and then ligated into a pcDNA3.1 vector (Invitrogen). The correct construction of the recombinant plasmids was verified through Sanger sequencing and extracted using an Endotoxin-Free Plasmid DNA Miniprep Kit (Tiangen, China) before using plasmids. The primers were listed in Supplementary file 1.

Protein Purification and Analysis

Earlier version: For protein purification, MDA5 plasmids with 6× His tag was constructed based on pcDNA3.1. (His6)-Flag-tagged MDA5 of M. miiuy and G. gallus were transiently overexpressed in 293T cells and lysed in a binding buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 20 mM Imidazole, and 1% Triton X-100) including protease inhibitor cocktail (Roche). The lysate was incubated over night at 4°C with anti-His beads (Solarbio). Anti-His beads were washed subsequently with washing buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 20 mM Imidazole, and 1% Triton X-100). MDA5 proteins were eluted by an addtion of elution buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 500 mM Imidazole, and 1% Triton X-100) to the beads. Purity of recombinant MDA5 was determined by SDS-PAGE separation.

Revised version: For protein purification, MDA5 plasmids with 6× His tag was constructed based on pcDNA3.1. (His6)-Flag-tagged MDA5 of M. miiuy and G. gallus were transiently overexpressed in 293T cells and lysed in a binding buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 20 mM Imidazole, and 1% Triton X-100) including protease inhibitor cocktail (Roche). The lysate was incubated over night at 4℃ with anti-His beads (Solarbio). Anti-His beads were washed subsequently with washing buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 20 mM Imidazole, and 1% Triton X-100). MDA5 proteins were eluted by an addtion of elution buffer (500 mM NaCl, 20 mM Tris/HCl [pH 7.4], 500 mM Imidazole, and 1% Triton X-100) to the beads. Purity of recombinant MDA5 was determined by SDS-PAGE separation.

RNA interference

The MDA5-specific small interfering RNA (si-MDA5-1 and si-MDA5-2) sequences were 5’-CGGACUACAUGCAGCGUAATT-3’ and 5’-GACCAAUGAGAUUUCUAUGTT-3’, respectively. The LGP2-specific small interfering RNA (si-LGP2) sequence was 5’-GGUUCACCUUGUAGAGCAATT-3’. The STING-specific small interfering RNA (si-STING) sequence was 5’-ACGACAGCAUCCAUUUCUATT-3’. The METTL3-specific small interfering RNA (si-METTL3-1, si-METTL3-2, and si-METTL3-3) sequences were 5’-CAUCACAAACGAACUCAACTT-3’, 5’-AACGUGGGCAAACUCUUUUTT-3’ and 5’-GAGAUUCUGGAACUACUUATT-3’, respectively. The METTL14-specific small interfering RNA (si-METTL14-1, si-METTL14-2, and si-METTL14-3) sequences were 5’-GCUGGAAAUCGAGGAGAUATT-3’, 5’-CCGTACGAAGAGGTGTACATT-3’ and 5’-GAGCCUCCCUUGGAAGAAUTT-3’, respectively. The YTHDF1-specific small interfering RNA (si-YTHDF1-1 and si-YTHDF1-2) sequences were 5’-AAAGACUUUGACUGGAACUTT-3’ and 5’-CAGUCGAUCAGAGACCUAATT-3’, respectively. The YTHDF2-specific small interfering RNA (si-YTHDF2-1 and si-YTHDF2-2) sequences were 5’-GCAGGGUGUUUAUCAUCAATT-3’, 5’-CUAUGCUCCCAGCUCAAUUTT-3’, respectively. The YTHDF3-specific small interfering RNA (si-YTHDF3-1 and si-YTHDF3-2) sequences were 5’-GGUGGACUACAAUGCCUAUTT-3’ and 5’-UCUACAGUAACAGCUAUGGTT-3’, respectively. The scrambled control RNA sequences were 5’ – UUCUCCGAACGUGUCACGUTT-3’. They were purchased from GenePharma (Shanghai, China).

Cells transfection

Transient transfection of cells with siRNA was performed in 24-well plates using Lipofectamine RNAiMAX (Invitrogen), and cells were transfected with DNA plasmids was performed using Lipofectamine 3000 (Invitrogen) according to the manufacturer’s instructions. For functional analyses, the overexpression plasmid (500 ng per well) or control vector (500 ng per well) and siRNA (100 nM) were transfected into cells in culture medium and then harvested for further detection.

RNA extraction and quantitative real-time PCR

Earlier version: Total RNA was extracted using the TRIzol Reagent (Invitrogen) and the cDNA was synthesized using the FastQuant RT Kit (Tiangen) which contains DNase treatment of RNA to eliminate genomic contamination, following the manufacturer’s instructions. The expression profiles of each gene were conducted by using the SYBR Premix Ex Taq™ (Takara), as previously described. The quantitative real-time PCR was conducted in an Applied Biosystems® QuantStudio 3 (Thermo Fisher Scientific). β-actin was used as internal controls for mRNA respectively. Primer sequences are listed in Supplementary file 1.

Revised version: Total RNA was extracted using the TRIzol Reagent (Invitrogen) and the cDNA was synthesized using the FastQuant RT Kit (Tiangen) which contains DNase treatment of RNA to eliminate genomic contamination, following the manufacturer’s instructions. The expression profiles of each gene were conducted by using the SYBR Premix Ex Taq™ (Takara), as previously described. The quantitative real-time PCR was conducted in an Applied Biosystems® QuantStudio 3 (Thermo Fisher Scientific). β-actin was used as internal controls. Primer sequences are listed in Supplementary file 1.

Western blotting

Cellular lysates were generated by using 1×SDS-PAGE loading buffer. Proteins were extracted from cells and measured with the BCA Protein Assay kit (Vazyme), then subjected to SDS-PAGE (10%) gel and transferred to PVDF (Millipore) membranes by semidry blotting (Bio-Rad Trans Blot Turbo System). The membranes were blocked with 5% BSA. Protein was blotted with different antibodies. The antibody against MDA5 was diluted at 1: 500 (Beyotime, Cat# AF7164); The antibody against IRF3 was diluted at 1: 500 (Boster, Cat# BA4351-2); anti-Flag, anti-Myc, anti-HA and anti-Tubulin monoclonal antibody were diluted at 1: 1,000 (Beyotime); and HRP-conjugated anti-rabbit IgG or anti-mouse IgG (Abbkine) at 1: 5,000. The results were the representative of three independent experiments. The immunoreactive proteins were detected by using WesternBrightTM ECL (Advansta). The digital imaging was performed with a cold CCD camera.

Dual-luciferase reporter assays

Earlier version: To determine the functional regulation of MDA5, MKC cells were co-transfected MDA5 expression plasmid, together with NF-κB, IRF3, IRF7, or IFN-1 luciferase reporter genes, and the pRL-TK Renilla luciferase reporter plasmid. After 48 h transfection, the cells were collected and lysed for reporter activity using the dual-luciferase reporter assay system (Promega). For the detection of m6A site, the wild or mutant of MDA5-exon1 luciferase reporter was co-transfected si-METTL3, si-METTL14, si-YTHDF1, si-YTHDF2, si-YTHDF3 or si-Ctrl into MKC cells. After 48 h transfection, the cells were collected and lysed for the reporter luciferase activities measured by using a dual-luciferase reporter assay system (Promega). All the luciferase activity values were gained against the Renilla luciferase control. For each experiment, three independent experiments were conducted, and each experiment was done in triplicate.

Revised version: For luciferase assays, MKC cells were co-transfected MDA5 expression plasmid, together with NF-κB, IRF3, IRF7, or IFN-1 luciferase reporter genes, and the pRL-TK Renilla luciferase reporter plasmid. For the detection of m6A site, the wild or mutant of MDA5-exon1 luciferase reporter was co-transfected si-METTL3, si-METTL14, si-YTHDF1, si-YTHDF2, si-YTHDF3 or si-Ctrl into MKC cells. After 48 h transfection, the cells were collected and lysed for the reporter luciferase activities measured by using a dual-luciferase reporter assay system (Promega). All the luciferase activity values were gained against the Renilla luciferase control.

RNA binding protein immunoprecipitation (RIP)

Earlier version: To identify whether MDA5 or MDA5-ΔRD can directly bind to SCRV viral RNA, MKC cells (∼2.0 × 107) were transfected with pcDNA3.1-Flag, pcDNA3.1-MDA5-Flag or pcDNA3.1-MDA5-ΔRD-Flag and then infected with SCRV at an MOI of 5 for 24 h. In addition, to identify whether YTHDF1, YTHDF2, and YTHDF3 can directly bind to MDA5, MKC cells were co-transfected with pcDNA3.1-Flag, pcDNA3.1-YTHDF1/2/3-Flag. For RNA immunoprecipitation (RIP) assays, MKC cells were harvested after 48 h transfection, and RIP assays were carried out with Magna RIP RNA-Binding Protein Immunoprecipitation Kit (Millipore) and anti-Flag antibody (Abcam) following the manufacturer’s protocol. Then, the expression of target gene SCRV-M, SCRV-G, β-actin, or MDA5 were detected by quantitative real-time qRT-PCR analysis.

Revised version: To identify whether MDA5 or MDA5-△RD can directly bind to SCRV viral RNA, MKC cells (∼2.0 × 107) were transfected with pcDNA3.1-Flag, pcDNA3.1-MDA5-Flag or pcDNA3.1-MDA5-△RD-Flag and then infected with SCRV at an MOI of 5 for 24 h. In addition, to identify whether YTHDF1, YTHDF2, and YTHDF3 can directly bind to MDA5, MKC cells were co-transfected with pcDNA3.1-Flag, pcDNA3.1-YTHDF1/2/3-Flag. For RNA immunoprecipitation (RIP) assays, MKC cells were harvested after 48 h transfection, and RIP assays were carried out with Magna RIP RNA-Binding Protein Immunoprecipitation Kit (Millipore) and anti-Flag antibody (Abcam) following the manufacturer’s protocol. Then, the expression of target gene SCRV-M, SCRV-G, β-actin, or MDA5 were detected by qRT-PCR analysis.

RNA Pull-Down assay

Earlier version: Cytoplasmic extracts were prepared from 2 × 107 MKC cells transfected with mmiMDA5-Flag (MDA5 of miiuy croaker) plasmids and 5 × 107 DF1 cells transfected with gga-MDA5-Flag (MDA5 of chicken) plasmids. The extracts were incubated with biotinylated 5’ppp-RNA or dsRNA and subjected to pull-down with streptavidin agarose beads (Geneseed), followed by subsequent Coomassie bluestaining or SDS-PAGE analysis and immunoblotting with anti-Flag antibody (Sigma).

Revised version: Cytoplasmic extracts were prepared from 2 × 107 MKC cells transfected with mmiMDA5-Flag (MDA5 of miiuy croaker) plasmids and 5 × 107 DF1 cells transfected with gga-MDA5-Flag (MDA5 of chicken) plasmids. The extracts were incubated with biotinylated 5’ppp-RNA or dsRNA and subjected to pull-down with streptavidin agarose beads (Geneseed), followed by subsequent Coomassie bluestaining or SDS-PAGE analysis and immunoblotting with anti-Flag antibody (Beyotime).

EMSA assay

Recombinant MDA5 proteins were mixed with biotin-labeled oligonucleotides in a reaction mixture (10 µl: 20 mM Tris-HCl pH 8.0, 1.5 mM MgCl2, 1.5 mM DTT). After incubation at room temperature for 15 min, the reaction mixtures were conducted on 6% non-denaturing polyacrylamide gels and the EMSA was performed using Chemiluminescent EMSA Kit (Beyotime).

MeRIP-seq assays and analysis

Earlier version: Total RNA was isolated using TRIzol reagent. m6A immunoprecipitation and library preparation were performed following the manufacturer’s protocol. Poly (A) RNA is purified from 50 μg total RNA using Dynabeads Oligo (dT) (Thermo Fisher) using two rounds of purification. Then the poly(A) RNA was fragmented into small pieces using Magnesium RNA Fragmentation Module (NEB) under 86°C for 7 min. Then the cleaved RNA fragments were incubated for 2 h at 4°C with m6A-specific antibody (Synaptic Systems) in IP buffer (50 mM Tris-HCl, 750 mM NaCl and 0.5% Igepal CA-630). Then the IP RNA was reverse-transcribed to create the cDNA by SuperScript™ II Reverse Transcriptase (Invitrogen). The average insert size for the final cDNA library was 300 ± 50 bp. For sequencing data analysis, we used HISAT2 to map reads to the reference genome of M. miiuy. Mapped reads of IP and input libraries were provided for R package exomePeak, which identifies m6A peaks with bed or bigwig format that can be adapted for visualization on the IGV software. MEME and HOMER were used for de novo and known motif finding followed by localization of the motif with respect to peak summit. Called peaks were annotated by intersection with gene architecture using R package ChIPseeker. The differentially expressed mRNAs were selected with log2 (fold change) >1 or log2 (fold change) <-1 and p-value < 0.05 by R package edgeR. The circus map of m6A peaks in untreated and SCRV-infected spleen tissues of M. miiuy was showed by TBtools35.

Revised version: Total RNA was isolated using TRIzol reagent. m6A immunoprecipitation and library preparation were performed following the manufacturer’s protocol. Poly (A) RNA is purified from 50 μg total RNA using Dynabeads Oligo (dT) (Thermo Fisher) using two rounds of purification. Then the poly(A) RNA was fragmented into small pieces using Magnesium RNA Fragmentation Module (NEB) under 86℃ for 7 min. Then the cleaved RNA fragments were incubated for 2 h at 4℃ with m6A-specific antibody (Synaptic Systems) in IP buffer (50 mM Tris-HCl, 750 mM NaCl and 0.5% Igepal CA-630). For sequencing data analysis, we used HISAT2 to map reads to the reference genome of M. miiuy. Mapped reads of IP and input libraries were provided for R package exomePeak, which identifies m6A peaks with bed or bigwig format that can be adapted for visualization on the IGV software. The differentially expressed mRNAs were selected with log2 (fold change) >1 or log2 (fold change) <-1 and p-value < 0.05 by R package edgeR. The circus map of m6A peaks in untreated and SCRV-infected spleen tissues of M. miiuy was showed by TBtools39.

MeRIP-qPCR

Earlier version: Total RNA was extracted using TRIzol reagent (Invitrogen). Methylated m6A RNA immunoprecipitation (Me-RIP) was performed according to the protocol of EpiQuikTM CUT&RUN m6A RNA Enrichment Kit (Epigentek). qPCR analysis of the methylated RNA was performed to detect methylated MDA5 mRNA levels.

Revised version: Total RNA was extracted using TRIzol reagent (Invitrogen). Methylated m6A RNA immunoprecipitation (Me-RIP) was performed according to the protocol of EpiQuikTM CUT&RUN m6A RNA Enrichment Kit (Epigentek). qPCR analysis of the methylated RNA was performed to detect methylated MDA5 mRNA levels. Relative m6A level for each transcript was calculated as the percent of input in each condition normalized to that of the respective positive control spike-in. Fold change of enrichment was calculated with mock samples normalized to 1 40.

Total m6A Quantification Assay

Total RNA was extracted by the TRIzol Reagent (Invitrogen). Then, an EpiQuik m6A RNA Methylation Quantification Kit (EpiGentek) was used to detect the total m6A level according to the manufacturer’s protocol. Briefly, positive control (PC), negative control (NC), and 200 ng isolated mRNA were added to each well with the capture antibody. Next, the RNA was incubated with m6A antibody in wells, and m6A content was quantified at a wavelength of 450 nm.

RNA stability assays

Cells were treated with actinomycin D (5 μg/mL) and then collected at different time points. RNA was extracted by TRIzol reagent (Invitrogen), and the mRNA levels were measured using qRT-PCR.

Database mining and sequence analysis

Earlier version: The species tree is constructed by submitting species names to the NCBI (https://www.ncbi.nlm.nih.gov/Taxonomy/CommonTree/wwwcmt.cgi). In order to determine whether there is RIG-I in vertebrates, we use H. sapiens and D. rerio RIG-I protein sequences to TBLASTN the whole genome sequence of the species on the ensemble website (http://www.ensemble.org/). In addition, in order to prevent the possibility of incomplete species genome, we also used RIG-I protein sequences of H. sapiens and D. rerio to compare the transcriptome and genome sequenced of the designated species on NCBI website (http://www.ncbi.nlm.nih.gov/Genbank/). For further gene synteny analysis, RIG-I of H. sapiens, C. milii, D. rerio were used as anchor sites. In order to identify the MDA5 and LGP2 from M. miiuy, we used the homologues of zebrafish reported previously as queries to seek for the transcriptome36 and a chromosome-scale genome database37 of M. miiuy by using local TBLASTN and BLASTN programs, then the protein structures were predicted by SMART website (http://smart.embl-heidelberg.de/).

Revised version: The species tree is constructed by submitting species names to the NCBI (https://www.ncbi.nlm.nih.gov/Taxonomy/CommonTree/wwwcmt.cgi). In order to determine whether there is RIG-I in vertebrates, we use H. sapiens and D. rerio RIG-I protein sequences to TBLASTN the whole genome sequence of the species on the ensemble website (http://www.ensemble.org/). In addition, in order to prevent the possibility of incomplete species genome, we also used RIG-I protein sequences of H. sapiens and D. rerio to compare the transcriptome and genome sequenced of the designated species on NCBI website (http://www.ncbi.nlm.nih.gov/Genbank/). For further gene synteny analysis, RIG-I of H. sapiens, C. milii, D. rerio were used as anchor sites. In order to identify the MDA5 and LGP2 from M. miiuy, we used the homologues of zebrafish reported previously as queries to seek for the transcriptome41 and a chromosome-scale genome database42 of M. miiuy by using local TBLASTN and BLASTN programs, then the protein structures were predicted by SMART website (http://smart.embl-heidelberg.de/).

Statistical analysis

Data are expressed as the mean ± SE from at least three independent triplicated experiments. A two-sided Student’s t test was used to evaluate the data. The relative gene expression data was acquired using the 2 -ΔΔCT method, and comparisons between groups were analyzed by one-way analysis of variance (ANOVA) followed by Duncan’s multiple comparison tests. A value of p < 0.05 was considered significant.

Supporting information